POMC neurons control fertility through differential signaling of MC4R in Kisspeptin neurons

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

The authors investigate the role of the melanocortin system in puberty onset. They conclude that POMC neurons within the arcuate nucleus of the hypothalamus provide important but differing input to kisspeptin neurons in the arcuate or rostral hypothalamus.

Strengths:

Innovative and novel

Technically sound

Well-designed

Thorough

Weaknesses:

There were no major weaknesses identified.

Reviewer #2 (Public review):

Summary:

This interesting manuscript describes a study investigating the role of MC4R signalling on kisspeptin neurons. The initial question is a good one. Infertility associated with MC4 mutations in humans has typically been ascribed to the consequent obesity and impaired metabolic regulation. Whether there is a direct role for MC4 in regulating the HPG axis has not been thoroughly examined. Here, the researchers have assembled an elegant combination of targetted loss of function and gain of function in vivo experiments, specifically targetting MC4 expression in kisspeptin neurons. This excellent experimental design should provide compelling evidence for whether melanocortin signalling dirently affects arcuate kisspeptin neurons to support normal reproductive function. There were definite effects on reproductive function (irregular estrous cycle, reduced magnitude of LH surge induced by exogenous estradiol). However, the magnitude of these responses and the overall effect on fertility were relatively minor. The mice lacking MC4R in kisspeptin neurons remained fertile despite these irregularities. The second part of the manuscript describes a series of electrophysiological studies evaluating the pharmacological effects of melanocortin signalling in kisspeptin cells in ex-vivo brain slides. These studies characterised interesting differential actions of melanocortins in two different populations of kisspeptin neurons. Collectively, the study provides some novel insights into how direct actions of melanocortin signalling via the MC4 receptor in kisspeptin neurons contribute to the metabolic regulation of the reproductive system. Importantly, however, it is clear that other mechanisms are also at play.

Strengths:

The loss of function/gain of function experiments provides a conceptually simple but hugely informative experimental design. This is the key strength of the current paper - especially the knock-in study that showed improved reproductive function even in the presence of ongoing obesity. This is a very convincing result that documents that reproductive deficits in MC4R knockout animals (and humans with deleterious MC4R gene variants) can be ascribed to impaired signalling in the hypothalamic kisspeptin neurons and not necessarily caused as a consequence of obesity. As concluded by the authors: "reproductive impairments observed in MC4R deficient mice, which replicate many of the conditions described in humans, are largely mediated by the direct action of melanocortins via MC4R on Kiss1 neurons and not to their obese phenotype." This is important, as it might change how such fertility problems are treated.

I would like to see the validation experiments for the genetic manipulation studies given greater prominence in the manuscript because they are critical to interpretation. Presently, only single unquantified images are shown, and a much more comprehensive analysis should be provided.

Weaknesses:

(1) Given that mice lacking MC4R in kisspeptin neurons remained fertile despite some reproductive irregularities, this can be described as a contributing pathway, but other mechanisms must also be involved in conveying metabolic information to the reproductive system. This is now appropriately covered in the discussion.

(2) The mechanistic studies evaluating melanocortin signalling in kisspeptin neurons were all completed in ovariectomised animals (with and without exogenous hormones) that do not experience cyclical hormone changes. Such cyclical changes are fundamental to how these neurons function in vivo and may dynamically alter how they respond to hormones and neuropeptides. Eliminating this variable makes interpretation difficult, but the authors have justified this as a reductionist approach to evaluate estradiol actions specifically. However, this does not reflect the actual complexity of reproductive function.

For example, the authors focus on a reduced LH response to exogenous estradiol in ovariectomised mice as evidence that there might be a sub-optimal preovulatory LH surge. However, the preovulatory LH sure (in intact animals) was not measured.

They have not assessed why some follicles ovulated, but most did not. They have focused on the possibility that the ovulation signal (LH surge) was insufficient rather than asking why some follicles responded and others did not. This suggests some issue with follicular development, likely due to changes in gonadotropin secretion during the cycle and not simply due to an insufficient LH surge.

Reviewer #3 (Public review):

The manuscript by Talbi R et al. generated transgenic mice to assess the reproduction function of MC4R in Kiss1 neurons in vivo and used electrophysiology to test how MC4R activation regulated Kiss1 neuronal firing in ARH and AVPV/PeN. This timely study is highly significant in neuroendocrinology research for the following reasons.

(1) The authors' findings are significant in the field of reproductive research. Despite the known presence of MC4R signaling in Kiss1 neurons, the exact mechanisms of how MC4R signaling regulates different Kiss1 neuronal populations in the context of sex hormone fluctuations are not entirely understood. The authors reported that knocking out Mc4r from Kiss1 neurons replicates the reproductive impairment of MC4RKO mice, and Mc4r expression in Kiss1 neurons in the MC4R null background partially restored the reproductive impairment. MC4R activation excites Kiss1 ARH neurons and inhibits Kiss1 AVPV/PeN neurons (except for elevated estradiol).

(2) Reproduction dysfunction is one of obesity comorbidities. MC4R loss-of-function mutations cause obesity phenotype and impaired reproduction. However, it is hard to determine the causality. The authors carefully measured the body weight of the different mouse models (Figure 1C, Figure 2A, Figure 3B). For example, the Kiss1-MC4RKO females showed no body weight difference at puberty onset. This clearly demonstrated the direct function of MC4R signaling in reproduction but was not a consequence of excessive adiposity.

(3) Gene expression findings in the "KNDy" system align with the reproduction phenotype.

(4) The electrophysiology results reported in this manuscript are innovative and provide more details of MC4R activation and Kiss1 neuronal activation.

Overall, the authors have presented sufficient background in a clear, logical, and organized structure, clearly stated the key question to be addressed, used the appropriate methodology, produced significant and innovative main findings, and made a justified conclusion.

Comments on revisions:

The authors have addressed my comments.

Recommendations for the authors:

The reviewers noted that they received comments in response to their concerns, and some improvements have been made to the manuscript. However, as described below, in some cases, a rebuttal was provided, but changes were not made to the manuscript. It is suggested that these issues be addressed to improve the quality of the manuscript.

We thank the reviewers and editor for the assessment of the manuscript and recommendations for its improvement. We have addressed the remaining comments from reviewer #2 below, and hope that they find our revisions satisfactory.

Reviewer #2 (Recommendations for the authors):

The manuscript convincingly shows that MC4R in kisspeptin-producing cells can influence reproductive function. This suggests that fertility problems associated with melanocortin mutations are likely due to direct effects on the reproductive systems rather than simply being side effects of the resultant obesity.

We are pleased that this reviewer finds the data convincing and thank them for the careful review of the manuscript, which has helped to improve its published version.

The authors have responded to the reviewer's comments and made several improvements to the manuscript.

The authors are correct in pointing out that the POMC-Cre animals should be fine for studies involving the administration of AAVs to adult animals. I have misinterpreted how these mice were being used, and this concern is fully addressed.

Unfortunately, in some cases, the authors rebutted the reviewer's comments but did not change the manuscript. I suggest addressing several issues in the manuscript (after all, it is not the reviewer's opinion that counts; this process is about improving the manuscript).

(1) Validation of the KO is insufficiently reported. From the methods, it appears that this was done thoroughly, but currently, only a single image of the arcuate nucleus is shown, and no image of the AVPV is shown. There is no quantitative information provided. The authors can keep these data as supplementary material, but they should be comprehensive and convincing, as so much depends on the degree of knockout in this model. One cannot assume complete KO based simply on the relevant genetics, as there are examples in this system where different Cre lines produce different outcomes with various floxed genes in the two major populations of kisspeptin neurons. This figure should show the quantitation of the RNAscope analysis from each of the two regions regarding the percentage of kisspeptin cells showing expression of MC4R mRNA. In addition, the lack of MC4 labelling in the arcuate nucleus, outside of kisspeptin neurons, is a concern. One would expect to see AgRP or POMC cells at this level, but are they still showing expression of MC4? A single image is insufficient to be convinced of the model's efficacy.

We appreciate the reviewer’s concerns regarding the validation of the MC4RKO model. Below, we provide clarification and additional justification for our approach.

(1) Quantification of MC4R in the Arcuate Nucleus (ARC): As noted by the reviewer, we were unable to detect sufficient MC4R signal in the ARC of KO mice to perform meaningful quantification. This is consistent with the expected outcome of a successful MC4R deletion. Given the low endogenous expression levels of MC4R in this region, even in control animals, and the technical limitations of RNAscope in detecting very low-abundance transcripts, especially for receptors, the absence of MC4R signal in the ARC of KO mice strongly supports effective deletion. Moreover, the MC4R loxP mouse has been published and validated by many labs including Brad Lowell’s lab who’s done extensive work using these mice for selective deletion of Mc4r from various neuronal populations such as Sim1 and Vglut2 neurons (Shah et al., 2014, de Souza Cordeiro et al., 2020). To further strengthen our validation, we provide additional images from another animal (Fig_S1) to illustrate the consistency of the MC4R KO in the ARC. These will be included as supplementary material, as suggested.Regarding AgRP and POMC neurons, MC4R is not highly expressed in these neurons (as per previous literature, e.g., Garfield et al., Nat Neurosci. 2015; Padilla SL et al, Endocrinology 2012; Henry et al, Nature, 2015). Instead, MC4R is predominantly found in downstream neurons in the paraventricular nucleus (PVN) and other hypothalamic regions (which is intact in our KO mice as shown in our validation figure). Thus, the absence of MC4R labeling in AgRP or POMC cells in our images aligns with known expression patterns and does not contradict the validity of our model.

(2) MC4R Expression in the AVPV and OVX Effect on Kiss1 Expression: We acknowledge the reviewer’s request for MC4R expression analysis in the anteroventral periventricular nucleus (AVPV). However, due to the timing of tissue collection after ovariectomy (OVX), Kiss1 expression in the AVPV is significantly suppressed, making it technically unfeasible to perform co-staining of MC4R with Kiss1 in this region. This is a well-documented effect of estrogen depletion following OVX (Smith et al., 2005; Lehman et al., 2010). While we acknowledge that an ideal validation would include AVPV co-labeling, the experimental constraints related to OVX preclude this analysis in our dataset.

Given these considerations and validations, we are confident that the KO is effective and specific.

(2) Line 88: "... however, conflicting reports exist". Expand on this sentence to describe what these conflicting reports show. The authors responded to my comment but made no changes to the introduction. As a reader, I dislike being told there are conflicting reports, but then I have to go and look up the reference to see what that actual point of conflict is.

By conflicting reports we meant that other studies have shown no association between MC4R and reproductive disorders, this has now been included in the revised manuscript (Line 89).

(3) Could the authors explain how a decrease in AgRP would be interpreted as a "decrease in hypothalamic melanocortin tone" in line 142 and line 364? These overly simplistic interpretations of qPCR data detract from the overall quality of the paper.

The reference to a decrease in melanocortin tone referred to the decrease in the expression of melanocortin receptor signaling, this has been clarified in the revised manuscript (lines 142 and 360).

(4) Please show the individual cycle patterns for all animals, as in Figure 2B. This can be a supplemental figure, but the current bar charts are not informative.

We respectfully disagree that the bar charts are not informative as they include the critical statistical analysis. We have now included all individual estrous cycle data in new separate supplemental figure (Sup. Figure 3). Therefore, we have excluded the representative cycles from the main figures as they are now in the new Supplemental. We have changed the orders of the figures in the text accordingly.

(5) In their rebuttal, the authors state: "Mice lack true follicular and luteal phases, and therefore, it is impossible to separate estrogen-mediated changes from progesterone-mediated changes (e.g., in a proestrous female). Therefore, we use an ovariectomized female model in which we can generate an LH surge with an E2-replacement regimen [1]. This model enables us to focus on estrogen effects, exclude progesterone effects, and minimize variability. Inclusion of cycling females would make interpretation much more difficult." I disagree, but the authors can take this position if they wish. However, they should not report the responses to exogenous estradiol in an ovariectomised mouse as a "preovulatory LH surge" (line 380). An ovariectomised mouse cannot ovulate, and the estrogen-induced LH surge is significantly different in magnitude and timing from the endogenous preovulatory LH surge (likely due to the actions of progesterone). One goal of these studies is to understand why the ovulation rate appears to be low in the MC4-KO animals. Hence, evaluating whether the preovulatory LH surge is typical is important. This has not been done. The authors have shown that the response to exogenous estradiol is sub-normal. Such an effect might lead to a reduced preovulatory LH surge, but this has not been measured.

We appreciate this reviewer’s concern about the nature of the preovulatory LH surge. We have clarified this in the revised manuscript and described it as “an induced LH surge” throughout the text (Lines 163, 533, 6560).

(6) I believe that the ovulation process should be considered "all or none," and I do not quite understand the rebuttal discussion. The authors describe that "numerous follicles mature at the same time....". That is not disputed. My point was that each mature follicle will receive the identical endocrine ovulatory signal (correct? Or do the authors believe something different?). If it were sufficient for one follicle to ovulate, then all of those mature follicles (the number of which will be variable between animals and between cycles) would be expected to undergo ovulation. The fact that they do not raise several possibilities. One that the authors favor is that an insufficient ovulatory signal might approach a threshold where some follicles ovulate and others do not. This possibility is supported by the apparent increase in cystic follicles, which might be preovulatory follicles that did not complete the ovulation process. Such variation might be stochastic, within normal variation for sensitivity to LH. However, it is also possible that the follicles have not matured at the same rate, perhaps influenced by abnormal secretion of LH or FSH during earlier phases of the cycle, and hence are not in the appropriate condition to respond to the ovulation signal when it arrives. Some may even have matured prematurely due to the elevated gonadotropins reported in this study. Given the data and the partial fertility, the most likely explanation is that the genetic manipulation has resulted in fewer follicles being available for ovulation due to changes in follicular development rather than a deficit of the ovulation signal, although the latter mechanism might also contribute. A third possibility is that genetic manipulation has directly affected the ovary. The authors did not answer whether Kiss1 and MC4 are co-expressed in the ovary. I think the authors might want to rule this out by showing no change in MC4R expression in the ovary.

We thank the reviewer for this thoughtful comment and agree that these are possible outcomes. We have now acknowledged them in the Discussion.

To answer the reviewer’s question, we have not investigated the co-expression of Kiss1 and Mc4r in the ovary. While MC4R has indeed been documented in the ovary (Chen et al. Reproduction, 2017), the changes in gonadotropin release and supporting in vitro data included in this manuscript clearly document a central effect, however, an additional effect at the level of the ovary cannot be completely ruled out. This has now been added to the discussion (Line 378-387).

(7) Lines 390, 454 " impaired LH pulse" What was the evidence for impaired LH pulse (see figure 2D)?

Thank you for pointing this out. This comment referred to augmented LH release. This has been corrected in the revised manuscript (Line 394).

The paper's strengths remain, as outlined in my original review. The authors have addressed what I perceived to be weaknesses, predominantly by changing the tone of discussion and interpretation of the data. This is appropriate. I consider the focus on the LH surge as the primary mechanism too narrow, and the authors should be considering how other changes during the cycle might influence ovarian function.

We sincerely appreciate the reviewer’s thoughtful evaluation of our manuscript and their constructive feedback. We are pleased that our revisions have addressed the perceived weaknesses and that the adjustments to the discussion and interpretation were deemed appropriate.

We acknowledge the reviewer’s perspective on broadening the discussion beyond the LH surge to consider additional cycle-dependent influences on ovarian function. While our current study focuses on this specific mechanism, we recognize that ovarian function is influenced by multiple physiological changes throughout the cycle. We have refined our discussion to reflect this broader context and appreciate the suggestion to consider these additional factors in future studies.

We have addressed all of the reviewer’s comments to the best of our ability and hope they find the revised manuscript satisfactory.

Changes from the earlier version to the revised version

Abstract

Abstract

Inactivating mutations in the melanocortin 4 receptor (MC4R) gene cause monogenic obesity. Interestingly, female patients also display various degrees of reproductive disorders, in line with the subfertile phenotype of MC4RKO female mice. However, the cellular mechanisms by which MC4R regulates reproduction are unknown. Kiss1 neurons directly stimulate gonadotropin-releasing hormone (GnRH) release through two distinct populations; the Kiss1ARH neurons, controlling GnRH pulses, and the sexually dimorphic Kiss1AVPV/PeN neurons controlling the preovulatory LH surge. Here, we show that Mc4r expressed in Kiss1 neurons regulates fertility in females. In vivo, deletion of Mc4r from Kiss1 neurons in female mice replicates the reproductive impairments of MC4RKO mice without inducing obesity. Conversely, reinsertion of Mc4r in Kiss1 neurons of MC4R null mice restores estrous cyclicity and LH pulsatility without reducing their obese phenotype. In vitro, we dissect the specific action of MC4R on Kiss1ARH vs Kiss1AVPV/PeN neurons and show that MC4R activation excites Kiss1ARH neurons through direct synaptic actions. In contrast, Kiss1AVPV/PeN neurons are normally inhibited by MC4R activation except under elevated estradiol levels, thus facilitating the activation of Kiss1AVPV/PeN neurons to induce the LH surge driving ovulation in females. Our findings demonstrate that POMCARH neurons acting through MC4R, directly regulate reproductive function in females by stimulating the “pulse generator” activity of Kiss1ARH neurons and restricting the activation of Kiss1AVPV/PeN neurons to the time of the estradiol-dependent LH surge, and thus unveil a novel pathway of the metabolic regulation of fertility by the melanocortin system.

Introduction

Obesity rates have skyrocketed in Western societies in the last decades, resulting in an alarming rise in comorbidities that place a significant burden on healthcare systems. The increase in obesity correlates with a decrease in fertility rates, leading to conception challenges that are experienced by approximately 15% of couples in the United States currently (1).

Earlier version: The melanocortin 4 receptor (MC4R) binds i) α-melanocyte stimulating hormone (αMSH), an agonist product of the pro-opiomelanocortin (Pomc) gene, and ii) the inverse agonist, agouti-related peptide (AgRP), to regulate food intake and energy expenditure (2, 3). While the role of MC4R on food intake is largely mediated by neurons located in the paraventricular nucleus of the hypothalamus (PVN) (4), its expression in the brain is widespread (5) with the specific function of the different MC4R-expressing neurons in areas beyond the PVN remaining to be fully explored. Inactivating mutations in MC4R are a leading cause of monogenic obesity in humans (6) and cause excessive obesity and hyperphagia in mice (7). Scant evidence in humans shows an association between MC4R mutations and higher incidence of hypogonadotropic hypogonadism (HH) (8), alterations in the timing of puberty onset (9), and polycystic ovary syndrome (PCOS) (10); however, conflicting reports exist (6). Therefore, the role of MC4R signaling in reproductive function in humans remains controversial despite the clear association found in mice. Supporting this role of MC4R, Mc4r null mice display an array of reproductive abnormalities that largely affects females, characterized by irregular estrous cycles, disrupted luteinizing hormone (LH) secretion, reduced corpora lutea and reduced fertility (11–13). Further evidence from mice demonstrates that MC4R agonists robustly increase LH release in a kisspeptin-dependent manner (14).

Revised version: The melanocortin 4 receptor (MC4R) binds i) α-melanocyte stimulating hormone (αMSH), an agonist product of the pro-opiomelanocortin (Pomc) gene, and ii) the inverse agonist, agouti-related peptide (AgRP), to regulate food intake and energy expenditure (2, 3). While the role of MC4R on food intake is largely mediated by neurons located in the paraventricular nucleus of the hypothalamus (PVN) (4), its expression in the brain is widespread (5) with the specific function of the different MC4R-expressing neurons in areas beyond the PVN remaining to be fully explored. Inactivating mutations in MC4R are a leading cause of monogenic obesity in humans (6) and cause excessive obesity and hyperphagia in mice (7). Scant evidence in humans shows an association between MC4R mutations and higher incidence of hypogonadotropic hypogonadism (HH) (8), alterations in the timing of puberty onset (9), and polycystic ovary syndrome (PCOS) (10); however, other studies have shown no association between MC4R and reproductive disorders (6). Therefore, the role of MC4R signaling in reproductive function in humans remains controversial despite the clear association found in mice. Supporting this role of MC4R, Mc4r null mice display an array of reproductive abnormalities that largely affects females, characterized by irregular estrous cycles, disrupted luteinizing hormone (LH) secretion, reduced corpora lutea and reduced fertility (11–13). Further evidence from mice demonstrates that MC4R agonists robustly increase LH release in a kisspeptin-dependent manner (14).

Kisspeptin is the most potent gonadotropin releasing hormone (GnRH) secretagogue known to date and it is mainly produced in two distinct neuronal populations. Kiss1 neurons of the arcuate nucleus of the hypothalamus (Kiss1ARH), present in both sexes, control the pulsatile (tonic) release of GnRH, and sex steroids attenuate their release of kisspeptin. Kiss1 neurons of the anteroventral periventricular continuum area (Kiss1AVPV/PeN), are predominantly present in females, and are responsible for generating the preovulatory GnRH/LH surge essential for ovulation (15). Despite these critical roles of both populations of Kiss1 neurons for reproduction, the underlying mechanisms that determine how each Kiss1 population responds differently to sex steroids to regulate the GnRH tonic vs surge release remains unresolved.

In rodents, compelling evidence indicates a close interaction between Kiss1 neurons and the melanocortin system: 1) fibers from POMC neurons in the arcuate nucleus (POMCARH) project to and juxtapose Kiss1ARH neurons (14); 2) the melanocortin signaling through MC4R contributes to the permissive role of leptin on puberty onset (14, 16, 17); and 3) MC4R expression on both Kiss1ARH (18–20) and Kiss1AVPV/PeN (18, 21). Altogether, this evidence suggests a clear role for MC4R in Kiss1 neurons, in the control of reproductive maturation and fertility.

In this study, we investigated the contribution of MC4R signaling versus obesity per se in the etiology of the reproductive impairments observed in Mc4r null mice, which could explain similar impairments observed in humans. Using genetic mouse models with specific deletion or reinsertion of Mc4r in Kiss1 neurons, we show that MC4R action in Kiss1 neurons is necessary for normal reproductive function in female mice. Whole cell voltage clamp recordings evidenced an excitatory action of MC4R on Kiss1ARH neurons and an estradiol-dependent inhibitory action on Kiss1AVPV/PeN neurons, with important implications for the timing of the preovulatory LH surge.

Results

Mc4r within Kiss1 neurons determines the timing of puberty onset in females

The expression of Mc4r within Kiss1ARH and Kiss1AVPV/PeN neurons has already been described elsewhere (18–21), suggesting a role for MC4R in the regulation of fertility. To further investigate the specific role of MC4R in Kiss1 neurons, we generated a Kiss1-specific MC4R knockout mouse model (Kiss1-MC4RKO). The specific deletion of Mc4r from Kiss1 neurons, as well as the absence of global recombination, were confirmed through RNAscope. While Kiss1 neurons of Kiss1-MC4RKO mice lack Mc4r transcript compared to their control littermates (Figure S1A), Mc4r was detectable in the PVN of all mice, which is a major site of Mc4r expression in the brain in the regulation of metabolism, therefore supporting the specific deletion of Mc4r only within Kiss1 neurons (Figure S1B).

Earlier version: Puberty onset was assessed daily from weaning age through the monitoring of vaginal opening (VO) and first estrus (FE). Kiss1-MC4RKO females showed a significant advancement in the age of VO (P=0.0150, Kiss1-MC4RKO: 24.67± 0.39 vs controls: 26.32±0.53) and FE (P=0.0341, Kiss1-MC4RKO: 32.63±1.4 vs controls: 37.84±1.90) compared to their littermate controls (Figure 1A, B). Body weight of Kiss1-MC4RKO females at the age of puberty onset was similar between groups (P=0.2596, Kiss1-MC4RKO: 12.41±0.30 vs controls: 12.89±0.26) (Figure 1C). These data suggest that melanocortin signaling on Kiss1 neurons participates in the timing of puberty onset in females. Since the role of the melanocortin system on puberty onset is largely unexplored, we evaluated the expression profile of the main components of this system in the hypothalami of WT female mice at postnatal days (PND) 10, 15, 22 and 30. Interestingly, expressions of Agrp, Mc3r and Mc4r were significantly lower at the time of puberty onset (PND30) compared to earlier developmental ages (Figure 1D). This supports our in vivo data and suggests that a decrease in the hypothalamic melanocortin tone drives puberty onset, in line with the advancement in the age of VO and FE observed in Kiss1-MC4RKO female mice (Figure 1A, B).

Revised version: Puberty onset was assessed daily from weaning age through the monitoring of vaginal opening (VO) and first estrus (FE). Kiss1-MC4RKO females showed a significant advancement in the age of VO (P=0.0150, Kiss1-MC4RKO: 24.67± 0.39 vs controls: 26.32±0.53) and FE (P=0.0341, Kiss1-MC4RKO: 32.63±1.4 vs controls: 37.84±1.90) compared to their littermate controls (Figure 1A, B). Body weight of Kiss1-MC4RKO females at the age of puberty onset was similar between groups (P=0.2596, Kiss1-MC4RKO: 12.41±0.30 vs controls: 12.89±0.26) (Figure 1C). These data suggest that melanocortin signaling on Kiss1 neurons participates in the timing of puberty onset in females. Since the role of the melanocortin system on puberty onset is largely unexplored, we evaluated the expression profile of the main components of this system in the hypothalami of WT female mice at postnatal days (PND) 10, 15, 22 and 30. Interestingly, expressions of Agrp, Mc3r and Mc4r were significantly lower at the time of puberty onset (PND30) compared to earlier developmental ages (Figure 1D). This supports our in vivo data and suggests that a decrease in the hypothalamic melanocortin signaling drives puberty onset, in line with the advancement in the age of VO and FE observed in Kiss1-MC4RKO female mice (Figure 1A, B).

Mc4r expressed in Kiss1 neurons determines the timing of puberty onset.

Kiss1-MC4RKO females display advanced puberty onset, assessed by daily monitoring of vaginal opening (A) and first estrus (B), as documented by cumulative percent and mean age of animals at vaginal opening (A) and first estrus (B) in Kiss1-MC4RKO females (n=24) compared to WT littermates (n=19). *p<0.05 by Student’s t-test. Data presented as the mean± SEM. (C) Female Kiss1-MC4RKO (n=19) have normal body weight at the time of puberty onset compared to their WT littermates (n=15). (D) Ontogeny expression of melanocortin genes (Agrp, Pomc, Mc4r and Mc3r) in the ARH of WT female mice at different postnatal pre-pubertal and pubertal ages: P10, P15, P22 and P30, normalized to the housekeeping gene Hprt. (n values: females at P10 (n=6), P15 (n=6), P22 (n=6), and P30 (n=5)). Groups with different letters are significantly different (p<0.05), as determined by one-way ANOVA followed by the Student-Newman-Keuls test. Data presented as the mean± SEM.

MC4R in Kiss1 neurons is required for female reproduction

Earlier version: Because Kiss1-MC4RKO females showed altered puberty onset, we further investigated their reproductive phenotype. Interestingly, Kiss1-MC4RKO female mice displayed normal BW throughout the time of the study (up to PND150) (Figure 2A); however, they presented with irregular estrous cycles with predominantly more time spent in diestrus (P=0.0001) and less time spent in estrus (P<0.0001) than their control littermates (Figure 2B, C). The assessment of the pulsatile release of LH every 10 min over 180 min revealed no change in the total number of pulses between groups but a significant increase in basal LH release (P=0.04) in Kiss1-MC4RKO female mice (0.42±0.03) compared to controls (0.34±0.02) (Figure 2D-H). The analysis of the gene expression of the “KNDy” systems in the ARH, which control the GnRH pulse generator (15), revealed normal expression levels of Kiss1, Tac2 and Tacr3, but significantly lower expression of Pdyn (P=0.0067, Kiss1-MC4RKO: 0.78±0.038 vs controls: 1.00±0.019) (Figure 2I-L). The conserved expression of Kiss1, Tac2 and Tacr3 correlates with the preserved LH pulse frequency and amplitude in Kiss1-MC4RKO mice, while the lower inhibitory tone of dynorphin (Figure 2I) correlates with the higher basal LH levels (Figure 2E). However, the mRNA expression of these genes, alone, does not necessarily reflect changes of their activities. To assess the contribution of MC4R signaling in Kiss1 neurons to the induction of ovulation through the preovulatory LH surge, Kiss1-MC4RKO females and control littermates were submitted to an LH surge induction protocol that showed the LH surge was significantly blunted in the Kiss1-MC4RKO females (P=0.0091), while the protocol clearly evoked the expected afternoon rise of LH in control mice (Figure 2M). In line with these findings, the ovaries of Kiss1-MC4RKO females displayed fewer corpora lutea, markers of recent ovulation (P=0.0054, Kiss1-MC4RKO: 0.80±0.37 vs controls: 2.80±0.37), in addition to increased cystic follicles (P=0.0337, Kiss1-MC4RKO: 1.60±0.40 vs controls: 0.40±0.24) (Figure 2N-P), which correlate with decreased fertility as observed by the extended time to deliver pups (i.e. longer time to get pregnant) (P=0.0030, Kiss1-MC4RKO: 30.30±4.60 vs controls: 20.88±0.22), and fewer pups per litter (P=0.0095, Kiss1-MC4RKO: 6.66±0.55 vs controls: 8.44±0.24) (Figure 2S, T).

Revised version: Because Kiss1-MC4RKO females showed altered puberty onset, we further investigated their reproductive phenotype. Interestingly, Kiss1-MC4RKO female mice displayed normal BW throughout the time of the study (up to PND150) (Figure 2A); however, they presented with irregular estrous cycles with predominantly more time spent in diestrus (P=0.0028) and less time spent in estrus (P=0.0067) than their control littermates (Figure 2B, S3a, b). The assessment of the pulsatile release of LH every 10 min over 180 min revealed no change in the total number of pulses between groups but a significant increase in basal LH release (P=0.04) in Kiss1-MC4RKO female mice (0.42±0.03) compared to controls (0.34±0.02) (Figure 2C-H). The analysis of the gene expression of the “KNDy” systems in the ARH, which control the GnRH pulse generator (15), revealed normal expression levels of Kiss1, Tac2 and Tacr3, but significantly lower expression of Pdyn (P=0.0067, Kiss1-MC4RKO: 0.78±0.038 vs controls: 1.00±0.019) (Figure 2I-L). The conserved expression of Kiss1, Tac2 and Tacr3 correlates with the preserved LH pulse frequency and amplitude in Kiss1-MC4RKO mice, while the lower inhibitory tone of dynorphin (Figure 2I) correlates with the higher basal LH levels (Figure 2G). However, the mRNA expression of these genes, alone, does not necessarily reflect changes of their activities. To assess the contribution of MC4R signaling in Kiss1 neurons to the induction of ovulation through the preovulatory LH surge, Kiss1-MC4RKO females and control littermates were submitted to an LH surge induction protocol that showed the induced LH surge was significantly blunted in the Kiss1-MC4RKO females (P=0.0091), while the protocol clearly evoked the expected afternoon rise of LH in control mice (Figure 2M). In line with these findings, the ovaries of Kiss1-MC4RKO females displayed fewer corpora lutea, markers of recent ovulation (P=0.0054, Kiss1-MC4RKO: 0.80±0.37 vs controls: 2.80±0.37), in addition to increased cystic follicles (P=0.0337, Kiss1-MC4RKO: 1.60±0.40 vs controls: 0.40±0.24) (Figure 2N-P), which correlate with decreased fertility as observed by the extended time to deliver pups (i.e. longer time to get pregnant) (P=0.0030, Kiss1-MC4RKO: 30.30±4.60 vs controls: 20.88±0.22), and fewer pups per litter (P=0.0095, Kiss1-MC4RKO: 6.66±0.55 vs controls: 8.44±0.24) (Figure 2S, T).

Deletion of MC4R from Kiss1 neurons impairs fertility in Kiss1-MC4RKO females.

Earlier version: (A) Female Kiss1-MC4RKO (n=19) have normal body weight from weaning and until post-natal day (PND) 150 compared to their WT littermates (n=15). (B) Representative examples of estrous cycles in Kiss1-MC4RKO and control females (n=5/group) assessed by daily vaginal cytology for 15 days. (C) Kiss1-MC4RKO females displayed irregular estrous cycles with a longer time in diestrus and a shorter time in estrus compared to control females. *** P < 0.001 Two-way ANOVA. (D) Pattern of LH pulsatility was analyzed in gonad intact Kiss1-MC4RKO (n=6) and control (n=7) females. # Represents LH pulses. (E) LH total secretory mass assessed by area under the curve (AUC), (F) LH pulse amplitude, (G) Basal LH and (H) Total number of pulses/180 min were analyzed. *p<0.05 by Student’s t-test. The expression of the KNDy genes Pdyn (I), Kiss1 (J), Tac2 (K) and Tacr3 (L), was assessed in the ARH of adult Kiss1-MC4RKO (n=6) and control (n=3) females. **p<0.01 by Student’s t-test. Data presented as the mean± SEM. (M) Kiss1-MC4RKO (n=6) and control (n=5) females were subjected to an LH surge induction protocol. LH samples were collected in the morning (AM [8 a.m.]) and evening (PM [7 p.m.]) after lights off. **P < 0.01, Two-way ANOVA followed by Bonferroni post hoc test. (N) Ovarian histology shows a decrease in the number of corpora lutea (CL, O) and increase in the number of cystic follicles (CF, P) of Kiss1-MC4RKO compared to controls (n=5/group). # represents cystic follicles. Serum levels of Testosterone (Q) and AMH (R) in adult gonad intact Kiss1-MC4RKO and control females (n=5/group). Student t test for unpaired samples. Data presented as the mean± SEM. Kiss1-MC4RKO females display impaired fertility characterized by increased time to deliver pups (S) and decreased number of pups per litter (T), (n=9/group). *p<0.05 and **p<0.01 by Student’s t-test. Data presented as the mean± SEM.

Revised version: (A) Female Kiss1-MC4RKO (n=19) have normal body weight from weaning and until post-natal day (PND) 150 compared to their WT littermates (n=15). (B) Kiss1-MC4RKO females displayed irregular estrous cycles with a longer time in diestrus and a shorter time in estrus compared to control females. ** P < 0.01 Two-way ANOVA. (C, D) Pattern of LH pulsatility was analyzed in gonad intact Kiss1-MC4RKO (n=6) and control (n=7) females. # Represents LH pulses. (E) LH total secretory mass assessed by area under the curve (AUC), (F) LH pulse amplitude, (G) Basal LH and (H) Total number of pulses /180 min were analyzed. *p<0.05 by Student’s t-test. The expression of the KNDy genes Pdyn (I), Kiss1 (J), Tac2 (K) and Tacr3 (L), was assessed in the ARH of adult Kiss1-MC4RKO (n=6) and control (n=3) females. **p<0.01 by Student’s t-test. Data presented as the mean± SEM. (M) Kiss1-MC4RKO (n=6) and control (n=5) females were subjected to an LH surge induction protocol. LH samples were collected in the morning (AM [8 a.m.]) and evening (PM [7 p.m.]) after lights off. **P < 0.01, Two-way ANOVA followed by Bonferroni post hoc test. (N) Ovarian histology shows a decrease in the number of corpora lutea (CL, O) and increase in the number of cystic follicles (CF, P) of Kiss1-MC4RKO compared to controls (n=5/group). # represents cystic follicles. Serum levels of Testosterone (Q) and AMH (R) in adult gonad intact Kiss1-MC4RKO and control females (n=5/group). Student t test for unpaired samples. Data presented as the mean± SEM. Kiss1-MC4RKO females display impaired fertility characterized by increased time to deliver pups (S) and decreased number of pups per litter (T), (n=9/group). *p<0.05 and **p<0.01 by Student’s t-test. Data presented as the mean± SEM.

The increase in serum LH levels and decreased ovulation observed in the Kiss1-MC4RKO females is reminiscent of polycystic ovary syndrome (PCOS) mouse models (22, 23). Thus, we investigated whether Kiss1-MC4RKO females display a PCOS-like phenotype. We analyzed circulating levels of testosterone (T) and anti-müllerian hormone (AMH), which are frequently elevated in PCOS models (24). The Kiss1-MC4RKO females expressed normal T and AMH levels compared to control mice in diestrus (Figure 2Q, R). Thus, we can exclude a PCOS-like reproductive phenotype mediated by the lack of melanocortin signaling on Kiss1 neurons.

Re-insertion of MC4R in Kiss1 neurons of MC4RKO mice improves reproductive function

Earlier version: Kiss1-MC4RKO females displayed reproductive abnormalities resembling those described in MC4RKO females (11–13). Thus, we hypothesized that the reproductive defects described for the MC4RKO mice would be, at least in part, due to the absence of MC4R signaling in Kiss1 neurons. To further investigate this hypothesis, we generated mice that do not express Mc4r anywhere (MC4R-LoxTB, i.e. MC4RKO) or that express Mc4r only in Kiss1 neurons (Kiss1-cre:MC4R-loxTB). The specific re-insertion of Mc4r within Kiss1 neurons in the Kiss1-cre:MC4R-loxTB mice was confirmed through RNAscope (Figure S2A) Mc4r expression was not detected in the PVN of these mice (Figure S2B), and it was only detected in Kiss1 neurons in the Kiss1-cre:MC4R-loxTB mice. Puberty onset was assessed daily from weaning age through the monitoring of vaginal opening (VO) and first estrus (FE). MC4R-LoxTB and Kiss1-cre:MC4R-loxTB female mice displayed normal timing of puberty onset compared to their control littermates, as assessed by VO (P=0.104, MC4R-LoxTB: 30.86±1.90 vs Kiss1-cre:MC4R-loxTB: 30.09±1.00 vs controls: 27.57±0.81), and FE (P=0.8472, MC4R-LoxTB: 35.14±2.46 vs Kiss1-cre:MC4R-loxTB: 34.55±1.64 vs controls: 35.80±0.84) (Figure 3A), despite displaying significantly higher body weight at the time of puberty onset (P=0.0001, MC4R-LoxTB: 15.49±0.50 vs Kiss1-cre:MC4R-loxTB: 15.54±0.40 vs Controls: 13.29±0.30) (Figure 3B). As expected, MC4R-LoxTB females (MC4RKO) displayed increased body weight (Figure 3C) and an array of reproductive impairments that included: irregular estrous cycles with significantly more days in diestrus (P=0.0016) and fewer days in estrus (P=0.0317) compared to controls (Figure 3D, E), significantly higher serum LH levels characterized by higher basal (P=0.0882, MC4R-LoxTB: 0.48±0.02 vs Controls: 0.37±0.06) and amplitude levels per LH pulse (P=0.0102, MC4R-LoxTB: 0.730±0.009 vs Controls: 0.58±0.07) (Figure 3F-J), and fewer corporal lutea (P=0.0366, MC4R-LoxTB: 0.80±0.80 vs Controls: 4.20±1.15) (Figure 3K, L), recapitulating the same reproductive phenotype observed in Kiss1-MC4RKO mice despite the differences in BW. Re-introduction of MC4R into Kiss1 neurons in Kiss1-cre:MC4R-loxTB mice completely recovered estrous cyclicity, which was similar to controls (diestrus: P=0.1932, proestrus: P=0.8262, estrus: P=0.0547, compared to controls) (Figure 3D, E) despite Kiss1-cre:MC4R-loxTB mice showing the same degree of obesity as MC4R-LoxTB mice (Figure 3C), indicating that obesity per se was not mediating the irregular estrous cycles in MC4RKO mice. As indicated above, MC4R-LoxTB mice display higher overall circulating LH levels, and this feature was also recovered by the re-insertion of MC4R in Kiss1 neurons of Kiss1-cre:MC4R-loxTB mice (basal LH: 0.32±0.04, LH amplitude: 0.48±0.04, LH pulses/180min: 2.40±0.67 and AUC: 70.62±7.15 compared to controls) (Figure 3G-J), further confirming that the melanocortin action on Kiss1 neurons is required for the normal control of gonadotropin release. Despite these significant improvements in reproductive function, Kiss1-cre:MC4R-loxTB mice presented fewer corpora lutea than controls (0.80±0.80), similar to MC4R-LoxTB mice, suggesting that an ovulatory impairment still persists (Figure 3K, L). While tonic LH release and estrous cyclicity are predominantly controlled by Kiss1ARH neurons, ovulation is mediated by the induction of the preovulatory LH surge by Kiss1AVPV/PeN neurons. It is possible that the metabolic signals derived from their obese phenotype, and/or the absence of the direct action of MC4R in GnRH neurons (16, 25) prevents the complete recovery of ovulation in Kiss1-cre:MC4R-loxTB mice. Both genetic models displayed significantly lower T (MC4R-LoxTB: 28.50±6.20 vs Kiss1-cre:MC4R-loxTB: 31.22±3.90 vs Controls: 62.84±16.53) and AMH (MC4R-LoxTB: 150.6±22.0 vs Kiss1-cre:MC4R-loxTB: 146.9±10.8 vs Controls: 321.8±15.7) levels than control mice in diestrus (Figure 3M, N). Therefore, we can exclude, once again, a PCOS-like reproductive phenotype mediated by the lack of melanocortin signaling.

Revised version: Kiss1-MC4RKO females displayed reproductive abnormalities resembling those described in MC4RKO females (11–13). Thus, we hypothesized that the reproductive defects described for the MC4RKO mice would be, at least in part, due to the absence of MC4R signaling in Kiss1 neurons. To further investigate this hypothesis, we generated mice that do not express Mc4r anywhere (MC4R-LoxTB, i.e. MC4RKO) or that express Mc4r only in Kiss1 neurons (Kiss1-cre:MC4R-loxTB). The specific re-insertion of Mc4r within Kiss1 neurons in the Kiss1-cre:MC4R-loxTB mice was confirmed through RNAscope (Figure S2A). Mc4r expression was not detected in the PVN of these mice (Figure S2B), and it was only detected in Kiss1 neurons in the Kiss1-cre:MC4R-loxTB mice. Puberty onset was assessed daily from weaning age through the monitoring of vaginal opening (VO) and first estrus (FE). MC4R-LoxTB and Kiss1-cre:MC4R-loxTB female mice displayed normal timing of puberty onset compared to their control littermates, as assessed by VO (P=0.104, MC4R-LoxTB: 30.86±1.90 vs Kiss1-cre:MC4R-loxTB: 30.09±1.00 vs controls: 27.57±0.81), and FE (P=0.8472, MC4R-LoxTB: 35.14±2.46 vs Kiss1-cre:MC4R-loxTB: 34.55±1.64 vs controls: 35.80±0.84) (Figure 3A, B), despite displaying significantly higher body weight at the time of puberty onset (P=0.0001, MC4R-LoxTB: 15.49±0.50 vs Kiss1-cre:MC4R-loxTB: 15.54±0.40 vs Controls: 13.29±0.30) (Figure 3C). As expected, MC4R-LoxTB females (MC4RKO) displayed increased body weight (Figure 3D) and an array of reproductive impairments that included: irregular estrous cycles with significantly more days in diestrus (P=0.0016) and fewer days in estrus (P=0.0317) compared to controls (Figure 3E, S3c, e), significantly higher AUC (MC4R-LoxTB: 109.6±2.39 vs Controls: 80.36±9.05) serum LH levels characterized by higher basal (MC4R-LoxTB: 0.48±0.02 vs Controls: 0.37±0.06) and amplitude levels per LH pulse (MC4R-LoxTB: 0.730±0.009 vs Controls: 0.58±0.07) (Figure 3F-J), and fewer corporal lutea (P=0.0366, MC4R-LoxTB: 0.80±0.80 vs Controls: 4.20±1.15) (Figure 3K, L), recapitulating the same reproductive phenotype observed in Kiss1-MC4RKO mice despite the differences in BW. Re-introduction of MC4R into Kiss1 neurons in Kiss1-cre:MC4R-loxTB mice completely recovered estrous cyclicity, which was similar to controls (diestrus: P=0.1932, proestrus: P=0.8262, estrus: P=0.0547, compared to controls) (Figure 3E, S3c, d) despite Kiss1-cre:MC4R-loxTB mice showing the same degree of obesity as MC4R-LoxTB mice (Figure 3D), indicating that obesity per se was not mediating the irregular estrous cycles in MC4RKO mice. As indicated above, MC4R-LoxTB mice display higher overall circulating LH levels, and this feature was also recovered by the re-insertion of MC4R in Kiss1 neurons of Kiss1-cre:MC4R-loxTB mice (basal LH: 0.32±0.04, LH amplitude: 0.48±0.04, LH pulses/180min: 2.40±0.67 and AUC: 70.62±7.15 compared to controls) (Figure 3G-J), further confirming that the melanocortin action on Kiss1 neurons is required for the normal control of gonadotropin release. Despite these significant improvements in reproductive function, Kiss1-cre:MC4R-loxTB mice presented fewer corpora lutea than controls (0.80±0.80), similar to MC4R-LoxTB mice, suggesting that an ovulatory impairment still persists (Figure 3K, L). While tonic LH release and estrous cyclicity are predominantly controlled by Kiss1ARH neurons, ovulation is mediated by the induction of the preovulatory LH surge by Kiss1AVPV/PeN neurons. It is possible that the metabolic signals derived from their obese phenotype, and/or the absence of the direct action of MC4R in GnRH neurons (16, 25) prevents the complete recovery of ovulation in Kiss1-cre:MC4R-loxTB mice. Both genetic models displayed significantly lower T (MC4R-LoxTB: 28.50±6.20 vs Kiss1-cre:MC4R-loxTB: 31.22±3.90 vs Controls: 62.84±16.53) and AMH (MC4R-LoxTB: 150.6±22.0 vs Kiss1-cre:MC4R-loxTB: 146.9±10.8 vs Controls: 321.8±15.7) levels than control mice in diestrus (Figure 3M, N). Therefore, we can exclude, once again, a PCOS-like reproductive phenotype mediated by the lack of melanocortin signaling.

Re-insertion of MC4R in Kiss1 neurons restores estrous cyclicity and LH pulsatility in Kiss1-cre:MC4R-LoxTB females.

Earlier version: (A) Kiss1-cre:MC4R-LoxTB (n=11) and MC4R-loxTB (MC4RKO) (n=7) displayed normal puberty onset as compared to their control littermates (n=14), as documented by cumulative percent and mean age of animals at vaginal opening and first estrus. (B) Kiss1-cre:MC4R-LoxTB and MC4R-loxTB both had significantly higher body weights at the time of puberty onset compared to their controls. **P < 0.01, *** P < 0.001 by one way ANOVA. (C) MC4R-loxTB (n=7) and Kiss1-cre:MC4R-LoxTB (n=11) females displayed significantly higher body weight than their littermates (n=10) from post-natal day 30 onwards, *p<0.05 and **P < 0.01. Data presented as the mean± (SEM). (D) Representative examples of estrous cycles of MC4R-loxTB (n=4), Kiss1-cre:MC4R-LoxTB (n=5) and control littermate mice, (n=4) assessed by daily monitoring of vaginal smears for 15 days. (D/M: diestrus/metestrus, P: proestrus, E: estrus). (E) MC4R-loxTB females displayed irregular estrous cycles, presenting longer time in diestrus and shorter time in estrus compared to control females, while Kiss1-cre:MC4R-LoxTB females displayed regular estrous cyclicity, similar to controls. *p<0.05, **P < 0.01, *** P < 0.001, ****P < 0.0001, Two-way ANOVA. Data presented as the mean± SEM. (F) Pattern of LH pulsatility analyzed in gonad intact MC4R-loxTB (n=4), Kiss1-cre:MC4R-LoxTB (n=5) and control littermates females (n=5). LH samples were collected every 10 min for 180 minutes; # represents LH pulses. (G) LH total secretory mass, (H) LH pulse amplitude, (I) basal LH, and (J) total number of pulses/180 min were assessed. *p<0.05, **P < 0.01 by one way ANOVA. (K) Representative samples of ovarian histology from MC4R-loxTB, Kiss1-cre:MC4R-LoxTB and control females (n=5/group); CL: Corpora lutea. Data are presented as the mean± SEM. (L) Ovarian histology showed a significant decrease in the number of corpora lutea in the MC4RKO and Kiss1-cre:MC4R-LoxTB compared to controls. Groups with different letters are significantly different. Serum levels of (M) Testosterone and (N) AMH in adult gonad intact females MC4R-loxTB, Kiss1-cre:MC4R-LoxTB and their control littermates. *p<0.05, ****P < 0.0001, One Way ANOVA. Data are presented as the mean± SEM.

Revised version: (A, B) Kiss1-cre:MC4R-LoxTB (n=11) and MC4R-loxTB (MC4RKO) (n=7) displayed normal puberty onset as compared to their control littermates (n=14), as documented by cumulative percent and mean age of animals at vaginal opening and first estrus. (C) Kiss1-cre:MC4R-LoxTB and MC4R-loxTB both had significantly higher body weights at the time of puberty onset compared to their controls. **P < 0.01, *** P < 0.001 by one way ANOVA. (D) MC4R-loxTB (n=7) and Kiss1-cre:MC4R-LoxTB (n=11) females displayed significantly higher body weight than their littermates (n=10) from post-natal day 30 onwards, *p<0.05 and **P < 0.01. Data presented as the mean± (SEM). (E) MC4R-loxTB females displayed irregular estrous cycles, presenting longer time in diestrus and shorter time in estrus compared to control females, while Kiss1-cre:MC4R-LoxTB females displayed regular estrous cyclicity, similar to controls. *p<0.05, **P < 0.01, *** P < 0.001, ****P < 0.0001, Two-way ANOVA. Data presented as the mean± SEM. (F) Pattern of LH pulsatility analyzed in gonad intact MC4R-loxTB (n=4), Kiss1-cre:MC4R-LoxTB (n=5) and control littermates females (n=5). LH samples were collected every 10 min for 180 minutes; # represents LH pulses. (G) LH total secretory mass, (H) LH pulse amplitude, (I) basal LH, and (J) total number of pulses /180 min were assessed. *p<0.05, **P < 0.01 by one way ANOVA. (K) Representative samples of ovarian histology from MC4R-loxTB, Kiss1-cre:MC4R-LoxTB and control females (n=5/group); CL: Corpora lutea. Data are presented as the mean± SEM. (L) Ovarian histology showed a significant decrease in the number of corpora lutea in the MC4RKO and Kiss1-cre:MC4R-LoxTB compared to controls. Groups with different letters are significantly different. Serum levels of (M) Testosterone and (N) AMH in adult gonad intact females MC4R-loxTB, Kiss1-cre:MC4R-LoxTB and their control littermates. *p<0.05, ****P < 0.0001, One Way ANOVA. Data are presented as the mean± SEM.

Kiss1ARH neurons are excited by MC4R agonists

Earlier version: Based on the expression of Mc4r in Kiss1 neurons and the reproductive impairment found after Mc4r deletion within Kiss1 neurons (Figure 2), we hypothesized that Kiss1ARH neurons are direct targets of melanocortins and would respond to MC4R activation. Initially we did whole-cell, voltage clamp recordings in Kiss1ARH neurons (Figure 4A). Kiss1-Cre x Ai32 or Kiss1-Cre AAV-injected mice underwent ovariectomies (OVX) and were given estradiol (E2) replacement (see Methods). We targeted fluorescent cells for recording which were isolated synaptically by bathing the slices in tetrodotoxin (TTX, 1 μM). Focal application of the high-affinity melanocortin receptor agonist melanotan II (MTII, ∼250 nM) elicited a small inward current in half of the isolated Kiss1ARH cells (10/21) (Figure 4B). Next, we examined if the E2 state affected MTII response by recording from Kiss1ARH neurons from OVX females and found two thirds (6/9 cells) responded, but there was no significant difference in the average current (Figure 4C). Finally, it must be noted that at this concentration MTII is a non-selective MCR agonist, so we followed up using a perfusion of the MC4R-selective agonist THIQ (100 nM) to determine if activation of MC4Rs alone was sufficient to invoke a postsynaptic response in Kiss1ARH neurons from OVX+E2 females. Similar to MTII, THIQ was able to elicit an excitatory inward current in Kiss1ARH neurons (4/5 cells, -7.4±0.5 pA), indicating MC4R activation is sufficient (Figure 4D).

Revised version: Based on the expression of Mc4r in Kiss1 neurons and the reproductive impairment found after Mc4r deletion within Kiss1 neurons (Figure 2), we hypothesized that Kiss1ARH neurons are direct targets of melanocortins and would respond to MC4R activation. Initially we did whole-cell, voltage clamp recordings in Kiss1ARH neurons (Figure 4A). Kiss1-Cre x Ai32 or Kiss1-Cre AAV-injected mice underwent ovariectomy (OVX) and were given estradiol (E2) replacement (see Methods). We targeted fluorescent cells for recording which were isolated synaptically by bathing the slices in tetrodotoxin (TTX, 1 μM). Focal application of the high-affinity melanocortin receptor agonist melanotan II (MTII, ∼250 nM) elicited a small inward current in half of the isolated Kiss1ARH cells (10/21) (Figure 4B). Next, we examined if the E2 state affected MTII response by recording from Kiss1ARH neurons from OVX females and found two thirds (6/9 cells) responded, but there was no significant difference in the average current (Figure 4C). Finally, it must be noted that at this concentration MTII is a non-selective MCR agonist, so we followed up using a perfusion of the MC4R-selective agonist THIQ (100 nM) to determine if activation of MC4Rs alone was sufficient to invoke a postsynaptic response in Kiss1ARH neurons from OVX+E2 females. Similar to MTII, THIQ was able to elicit an excitatory inward current in Kiss1ARH neurons (4/5 cells, -7.4±0.5 pA), indicating MC4R activation is sufficient (Figure 4D).

Kiss1ARH neurons are activated by melanocortin agonists and respond to optogenetic stimulation of POMC neurons.

(A) Slices were taken from Kiss1-cre injected (confocal image) brains to target fluorescent cells (arrowhead) for recording (white outline over electrode). (B) Whole cell voltage clamp recording of Kiss1ARH neurons following the direct application of the high affinity melanocortin receptor agonist melanotan II (MTII, 250 nM) was added directly to the bath and the response was compared between slices from OVX and OVX+E females. (C) While the average inward current was slightly higher in the OVX+E state, this was not a significant difference (Student’s t-test for unpaired samples, p>0.05). (D) The MC4R-selective agonist THIQ (100 nM) was perfused, and excitatory inward currents were generated in Kiss1ARH neurons. (E) Images are of AAV-driven labeling of POMC cells as seen through the confocal. (F) Using the AAV-driven expression of channelrhodopsin in adult Pomc-cre mice, high frequency stimulation elicited a slow inward current in Kiss1 neurons. (G) The identity of cells was confirmed through the presence of a persistent sodium current (see (26)) and/or with RT-PCR of harvested cytoplasm showing Kiss1 expression. (Gel: MM=molecular marker, TC=tissue controls). (H) A direct synaptic projection from POMC to Kiss1ARH neuron (posthoc identified) was confirmed using the “rescue” protocol: 1, baseline glutamatergic responses were initially generated (black trace); 2, then action potentials were eliminated by blocking voltage-gated sodium channels with TTX and the postsynaptic response (red trace); 3, blockade of potassium channels facilitated calcium entry into the terminal through ChR2 to release synaptic vesicles, “rescuing” the postsynaptic glutamate response (green trace). I) Glutamatergic responses were also often observed, particularly in the ventral ARH. As seen when targeting low input resistance neurons (i.e., Kiss1ARH) with low frequency optogenetic stimulation (5 ms pulse, 50 ms inter-spike interval), the first response in OVX + vehicle females was larger relative to the second response in OVX + E2-treated female mice. Representative traces are the average of 30 sweeps. (J) The averaged paired-pulse ratio was lower in recordings from estradiol-treated female mice, indicating an increased release probability (Students t-test p<0.05,). (K) High frequency optogenetic stimulation elicited a small inward current. (L) In a different cell, a high frequency inward current was noted before constant perfusion of SHU9119 for 15 minutes (break between traces). When the stimulation protocol was repeated, no inward current was elicited. (M) The majority of high frequency responses were inward and only twice was an inhibitory outward current recorded in identified Kiss1ARH neurons. Perfusion of the non-selective opioid receptor antagonist naloxone reversed the current, eliminating MCR-activation as the mechanism.

POMC neurons synapse directly with Kiss1ARH neurons

Earlier version: Based on the pharmacological activation of MCRs, we wanted to address whether POMCARH neurons are the source of melanocortins to excite Kiss1ARH neurons directly. To answer this question, we injected an AAV-EF1α-DIO-ChR2:mCherry vector into the ARH of adult Pomc-Cre female mice (Figure 4E) (24). After 2-4 weeks we did whole cell recordings from putative Kiss1ARH neurons and looked for a response to high frequency optogenetic stimulation of POMC fibers in OVX + E2-treated females (Figure 4F). We were able to confirm that these were Kiss1 neurons based on the presence of a persistent sodium current and/or post hoc identification by scRT-PCR (40/77) cells, Figure 4G). While this current is more prevalent in the AVPV Kiss1 population, Kiss1 neurons are the only ARH neurons to display this electrophysiological “fingerprint” of a pronounced INaP paired with a high capacitance and a low input resistance (26). Additionally, while recording from putative Kiss1ARH neurons, we optogenetically stimulated POMC fibers at low frequency (1 Hz, 5 ms, 470 nm light) and recorded fast postsynaptic inward currents. The fast kinetics of the EPSC are also a tell-tale sign of an ionotropic glutamatergic response (27). Previously, CNQX was sufficient to block similar excitatory ESPCs in other postsynaptic targets of POMC neurons (28). While the consistent latency from optogenetic stimulus and current response was indicative of a single synaptic delay, we wanted to establish that it was a direct synaptic connection, as we have previously shown for output of POMC neurons (28) and Kiss1 neurons (29–31). First, we abrogated the optogenetic response with the addition of TTX (1 μM) to the bath, we were able to rescue the postsynaptic glutamate response with the addition of the K+ channel blockers 4-aminopyridine (4-AP; 0.5 mM) and tetraethyl ammonium (TEA; 7.5 mM) (Figure 4H). K+ channel block enables the calcium influx from ChR2 activation in the presynaptic terminals to be sufficient to restore vesicle fusion. Therefore, there appears to be no intervening synapses between POMC and the downstream Kiss1ARH neurons.

Revised version: Based on the pharmacological activation of MCRs, we wanted to address whether POMCARH neurons are the source of melanocortins to excite Kiss1ARH neurons directly. To answer this question, we injected an AAV-EF1α-DIO-ChR2:mCherry vector into the ARH of adult Pomc-Cre female mice (Figure 4E) (24). After 2-4 weeks we did whole cell recordings from putative Kiss1ARH neurons and looked for a response to high frequency optogenetic stimulation of POMC fibers in OVX + E2-treated females (Figure 4F). We were able to confirm that these were Kiss1 neurons based on the presence of a persistent sodium current and/or post hoc identification by scRT-PCR (40/77) cells, Figure 4G). While this current is more prevalent in the AVPV Kiss1 population, Kiss1 neurons are the only ARH neurons to display this electrophysiological “fingerprint” of a pronounced INaP paired with a high capacitance and a low input resistance (26). Additionally, while recording from putative Kiss1ARH neurons, we optogenetically stimulated POMC fibers at low frequency (1 Hz, 5 ms, 470 nm light) and recorded fast postsynaptic inward currents. The fast kinetics of the EPSC are also a tell-tale sign of an ionotropic glutamatergic response (27). Previously, CNQX was sufficient to block similar excitatory ESPCs in other postsynaptic targets of POMC neurons (28). While the consistent latency from optogenetic stimulus and current response was indicative of a single synaptic delay, we wanted to establish that it was a direct synaptic connection, as we have previously shown for output of POMC neurons (28) and Kiss1 neurons (29–31). First, we abrogated the optogenetic response with the addition of TTX (1 μM) to the bath, and then we were able to rescue the postsynaptic glutamate response with the addition of the K+ channel blockers 4-aminopyridine (4-AP; 0.5 mM) and tetraethyl ammonium (TEA; 7.5 mM) (Figure 4H). K+ channel block enables the calcium influx from ChR2 activation in the presynaptic terminals to be sufficient to restore vesicle fusion. Therefore, there appears to be no intervening synapses between POMC and the downstream Kiss1ARH neurons.

Previously we found that E2 increases glutamate release from POMC neurons by increasing Vglut2 (28). Therefore, we compared the postsynaptic glutamatergic responses following two optogenetic stimuli (50 ms interval between light flashes) in Kiss1ARH neurons from OVX + vehicle versus OVX + E2-treated female mice (Figure 4I). As anticipated, we found that treatment with E2 increased the probability of glutamate release from POMC neurons onto Kiss1ARH neurons based on the significant decrease in the paired pulse ratio of the two stimuli (32) (Figure 4J). Although in the present study we used an in vivo treatment paradigm, we know from previous studies that this augmentation of glutamate release can happen quite rapidly after a brief exposure to E2 in vitro (within 15 min) (28). Therefore, Kiss1ARH neurons are excited by the glutamatergic input from POMC neurons in an E2-dependent manner.

Earlier version: Next, we tested for a melanocortin response using a high frequency stimulation protocol known to elicit peptide release (33). Only cells displaying a glutamatergic response were tested. Indeed, high frequency stimulation evoked an excitatory inward current (Figure 4K) that was blocked by the selective MC3/4R antagonist SHU 9119 (Figure 4L). However, the number of high frequency responses seen in Kiss1 neurons was unexpectedly infrequent and small (4/16, 25%, mean inward current: -1.4±3.1 pA) compared to the efficacy of pharmacological MCR activation. E2-treatment enhances Pomc expression and β-endorphin labeling (34) (35), but the effect, if any, on αMSH release or the MCRs in Kiss1 neurons is unknown. We suspected that the high frequency stimulation was eliciting co-release of αMSH and β-endorphin which would exert oppositive effects on a postsynaptic cell. In order to isolate melanocortin signaling we pretreated slices with the non-selective opioid antagonist naloxone (1 μM) and found an increase in the likelihood and magnitude of a response to high frequency stimulation (9/23, 40%, mean inward current: -7.6±2.9 pA) (Figure 4M).

Revised version: Next, we tested for a melanocortin response using a high frequency stimulation protocol known to elicit peptide release (33). Only cells displaying a glutamatergic response were tested. Indeed, high frequency stimulation evoked an excitatory inward current (Figure 4K) that was blocked by the selective MC3/4R antagonist SHU 9119 (Figure 4L). However, the number of high frequency responses seen in Kiss1 neurons was unexpectedly infrequent and relatively small (4/16, 25%, mean inward current: -1.4±3.1 pA) compared to the efficacy of pharmacological MCR activation. E2-treatment enhances Pomc expression and β-endorphin labeling (34) (35), but the effect, if any, on αMSH release or the MCRs in Kiss1 neurons is unknown. We suspected that the high frequency stimulation was eliciting co-release of αMSH and β-endorphin, which would exert oppositive effects on a postsynaptic cell. In order to isolate melanocortin signaling we pretreated slices with the non-selective opioid antagonist naloxone (1 μM) and found an increase in the likelihood and magnitude of a response to high frequency stimulation (9/23, 40%, mean inward current: -7.6±2.9 pA) (Figure 4M).

Together these optogenetic findings reinforced our pharmacological results showing that Kiss1ARH neurons are excited by MCR agonists. Additional studies would need to be done to identify the cation conductance that is affected by the MC4R signaling cascade in Kiss1ARH neurons, but clearly we have established that there are excitatory glutamatergic and peptidergic inputs from POMC to Kiss1ARH neurons.

Kiss1AVPV/PeN neurons are inhibited by MC4R agonists

Having established a direct, excitatory projection from POMC to Kiss1ARH neurons, we next examined POMC inputs to Kiss1AVPVPeN neurons. First, we used immunocytochemistry to label fibers expressing αMSH to demonstrate that POMC fibers densely innervated the AVPV/PeN area (Figure 5A). Next, we recorded from AVPV neurons in slices taken from OVX+E2, POMC-Cre mice injected with AAV1-DIO-YFP:ChR2 into the ARH (Figure 5B). We targeted cells along the ventricle in the AVPV and PeN, surrounded by YFP terminals, and Kiss1AVPV/PeN neurons were identified based on the expression of several endogenous conductances (INaP, T-type calcium current, h-current) that are unique to these neurons (26, 36). We recorded a direct glutamatergic synaptic response following optogenetic stimulation in six neurons, which was further verified through pharmacological “rescue” of the synaptic response in the majority (5/6) neurons tested (Figure 5C), indicating that POMC neurons make direct monosynaptic connections with Kiss1AVPV/PeN neurons.

Kiss1AVPV/PeN neurons are inhibited by MC4R agonists in an E2-dependent manner.

(A) Immunohistochemistry showing robust labeling of α-MSH fibers in the AVPV/PeN region in OVX+E2 WT female mice. (B) Low and high power bright-field images taken during electrophysiology recording in OVX+E2 POMC-Cre mice expressing YFP:ChR2 in the ARH following AAV injection. (left) Low power image shows the location of recorded cells. (right) Higher power image of area from white box inset displaying POMC fibers innervating the area. (C) Fibers surrounded putative Kiss1-AVPV neurons (expressing INaP, IT and Ih) and optogenetic stimuli was able to elicit postsynaptic currents that were eliminated with tetrodotoxin (1 μM), but “rescued” with addition of K+ channel blockers (4-AP and TEA) which indicates a monosynaptic connection between POMCARH and Kiss1AVPV neurons. (D) Whole-cell voltage clamp recordings were made in ChR2-YFP positive cells in brain slices taken from Kiss1-cre:Ai32 female mice. Bath application of melanotan II (MTII, 500 nM) generated an inhibitory outward current in a Kiss1AVPV/PeN neuron from OVX female. As the cell was synaptically isolated using bath-applied tetrodotoxin (TTX, 1 μM) this represents a direct effect. Washout of MTII while still in TTX quickly led to a return to baseline RMP. (E) An IV relationship was plotted using voltage steps before and after MTII administration. The crossing at -80 mV (∼EK+) indicates that the opening of K+ channels underlie the MC4R inhibition of Kiss1AVPV neurons. (F) In a subset of recordings from OVX brains slices, the selective membrane estrogen receptor (Gq-mER) agonist STX (10 nM) was added to the bath for ∼10 minutes prior to addition of MTII. STX pretreatment resulted in either a greatly attenuated outward current or even an inward current as shown in this example. (G) The mean outward current was significantly higher when recording in Kiss1AVPV/Pen neurons from brain slices from vehicle-treated, OVX females compared to E2-treated, OVX females or acute STX-treated brain slices from OVX females (One-way ANOVA F(2,19)=12.32, p<0.001; Holm-Sidak posthoc comparisons found significant differences between all groups; * p<0.05, *** p<0.001). (H) There was no difference in the frequency of mEPSCs after MTII, calculated as a percent of baseline between groups. (I) However, there was a significant difference in the mEPSC amplitude (calculated as a percent of baseline) between all groups: main effect F(2,19) =10.58, p<0.001. Posthoc comparisons using Holm-Sidak found OVX+STX to be different from both OVX (p<0.001) and OVX+E (p<0.05).

We next hypothesized that Kiss1AVPV/PeN neurons would be excited by exogenous application of melanocortin agonists because the receptor is typically Gs-coupled (37, 38). We did whole-cell, voltage clamp recordings from Kiss1AVPV/PeN-Cre::Ai32 neurons, which were isolated synaptically by bathing the neurons in TTX (1 μM), from OVX females. For these experiments, we tested the response to the high affinity melanocortin receptor agonist MTII (500 nM). Surprisingly, we consistently measured an outward (inhibitory) current that could be reversed on washout of MTII (Figure 5D). This outward current was associated with an increase in a K+ conductance based on the current-voltage plot (i.e., the outward current reversed at ∼EK+, Figure 5E). Also, the inwardly-rectifying I/V could indicate that MC4R is coupled to activation of inwardly-rectifying K+ channels as has been previously reported for MC4R signaling in pre-autonomic parasympathetic neurons in the brainstem (39). We hypothesized that the coupling could be modulated by E2 via a Gαq-coupled membrane estrogen receptor (Gq-mER) as we have previously demonstrated in POMC neurons (40–42). For these experiments, we again utilized OVX females and targeted Kiss1AVPV/PeN-Cre::Ai32 neurons. Once in the whole-cell voltage-clamp configuration, we perfused the slices with STX (10 nM), a selective ligand for the putative Gq-mER (42), and tested the response to the MTII. Indeed, the outward current (inhibitory) response to MTII was completely abrogated and even reversed by STX (Figure 5F, G). The short-term (bath) treatment with STX ensured that there was no desensitization of the Gq-mER with longer-term (in vivo) treatment with E2. Therefore, estrogen receptor activation can rapidly uncouple (i.e., desensitize) the MC4R inhibitory response in Kiss1AVPV/PeN neurons. We would predict that the intracellular signaling cascade for the heterologous desensitization is similar to what we have elucidated in POMC neurons (42), but this will need to be determined in future experiments. We next investigated the response in vehicle-treated, OVX females. Indeed, Kiss1AVPV/PeN neurons were even more inhibited by the same MTII exposure (Figure 5G). Therefore, in contrast to Kiss1ARH neurons, the MC4R appears to be coupled to opening of K+ channels in Kiss1AVPV/PeN neurons. To further establish whether the MTII was having pre- or postsynaptic effects, we measured the glutamatergic mEPSCs before and after melanocortin receptor activation. We found there was no significant effect on the frequency (Figure 5H), but there was an effect on the amplitude (Figure 5I). This further supports a postsynaptic locus of STX’s effects. However, posthoc comparisons found STX increased amplitude compared to both OVX and OVX+E. This result hints at a membrane delimited effect as STX is a much more potent agonist for Gq-mER than E2 (43).

Discussion

Earlier version: Our findings show that melanocortins act directly on Kiss1 neurons to contribute to the metabolic regulation of fertility, in line with previous reports showing that αMSH stimulates LH release in a kisspeptin-dependent manner (14) and that Kiss1 neurons express the melanocortin receptor MC4R (18–21). Several studies in humans have linked Mc4r mutations to reproductive abnormalities, including precocious puberty (9), PCOS (10), and hypogonadism (8). However, a direct association between MC4R mutations and reproductive function, independent from the obese condition of these patients, has not been identified (6). In the current study, we show that the reproductive impairments observed in MC4R deficient mice, which replicate many of the conditions described in humans, are largely mediated by the direct action of melanocortins via MC4R on Kiss1 neurons and not to their obese phenotype. This is because the ablation of MC4R from Kiss1 neurons largely replicated the reproductive impairments observed in MC4RKO female mice without inducing obesity, and the selective reinsertion of MC4R into Kiss1 neurons of MC4RKO mice significantly improved their reproductive function without changing their obese phenotype. Strikingly, puberty onset was advanced in Kiss1-MC4RKO females. Our findings revealed a low melanocortin tone in the hypothalamus of WT females during pubertal development, characterized by lower levels of Agrp, Mc4r and Mc3r expression at the time of puberty onset (∼PND30). This developmental decline in the expression of melanocortin receptors is in line with the advanced timing of puberty onset in females when MC4R is congenitally ablated from Kiss1 neurons. A recent study using whole body MC3RKO mice showed a trend to delayed puberty onset (19), suggesting a larger role of melanocortins in the timing of puberty onset. Here, we show that in females the combination of MC3R and MC4R action may be necessary for pubertal development. Interestingly, the MC4R-LoxTB and Kiss1-cre:MC4R-LoxTB mice do not show an advancement in puberty onset as we observed in Kiss1-MC4RKO mice; however, these mice are obese at the time of puberty onset already, which may affect the timing of sexual maturation independently of MC4R action in Kiss1 neurons. Of note, the findings presented in this study do not rule out the existence of additional sites of action of MC4R in other neuronal populations to regulate reproduction, e.g., GnRH neurons, which also express MC4R (16, 25). Indeed, the lack of full recovery of the reproductive function in Kiss1-cre:MC4R-LoxTB mice suggests that MC4R in Kiss1 neurons is necessary but not sufficient to achieve full reproductive capabilities.

Revised version: Our findings show that melanocortins act directly on Kiss1 neurons to contribute to the metabolic regulation of fertility, in line with previous reports showing that αMSH stimulates LH release in a kisspeptin-dependent manner (14) and that Kiss1 neurons express the melanocortin receptor MC4R (18–21). Several studies in humans have linked Mc4r mutations to reproductive abnormalities, including precocious puberty (9), PCOS (10), and hypogonadism (8). However, a direct association between MC4R mutations and reproductive function, independent from the obese condition of these patients, has not been identified (6). In the current study, we show that the reproductive impairments observed in MC4R deficient mice, which replicate many of the conditions described in humans, are largely mediated by the direct action of melanocortins via MC4R on Kiss1 neurons and not to their obese phenotype. This is because the ablation of MC4R from Kiss1 neurons largely replicated the reproductive impairments observed in MC4RKO female mice without inducing obesity, and the selective reinsertion of MC4R into Kiss1 neurons of MC4RKO mice significantly improved their reproductive function without changing their obese phenotype. Strikingly, puberty onset was advanced in Kiss1-MC4RKO females. Our findings revealed a low melanocortin expression in the hypothalamus of WT females during pubertal development, characterized by lower levels of Agrp, Mc4r and Mc3r expression at the time of puberty onset (∼PND30). This developmental decline in the expression of melanocortin receptors is in line with the advanced timing of puberty onset in females when MC4R is congenitally ablated from Kiss1 neurons. A recent study using whole body MC3RKO mice showed a trend to delayed puberty onset (19), suggesting a larger role of melanocortins in the timing of puberty onset. Here, we show that in females the combination of MC3R and MC4R action may be necessary for pubertal development. Interestingly, the MC4R-LoxTB and Kiss1-cre:MC4R-LoxTB mice do not show an advancement in puberty onset as we observed in Kiss1-MC4RKO mice; however, these mice are obese at the time of puberty onset already, which may affect the timing of sexual maturation independently of MC4R action in Kiss1 neurons. Of note, the findings presented in this study do not rule out the existence of additional sites of action of MC4R in other neuronal populations to regulate reproduction, e.g., GnRH neurons, which also express MC4R (16, 25). Indeed, the lack of full recovery of the reproductive function in Kiss1-cre:MC4R-LoxTB mice suggests that MC4R in Kiss1 neurons is necessary but not sufficient to achieve full reproductive capabilities.

Earlier version: In addition to early puberty onset, Kiss1-MC4RKO females displayed an impaired preovulatory LH surge driving ovulation and had fewer corpora lutea in their ovaries compared to control littermates, further supporting a role of MC4R in regulating ovulation. In adulthood, the reproductive phenotype observed in Kiss1-MC4RKO and MC4R-loxTB female mice (increased LH, irregular estrous cycles, oligo-ovulation, increased cystic follicles) correlates with the phenotype observed in PCOS mouse models (22, 23). In fact, an association between MC4R mutations and PCOS has been reported (10). However, our mouse models failed to display higher levels of circulating androgens or AMH, two of the hallmarks of PCOS (24), suggesting that melanocortin signaling is unlikely to be a main contributing factor in the development of this syndrome.

Revised version: In addition to early puberty onset, Kiss1-MC4RKO females displayed an impaired preovulatory LH surge driving ovulation and had fewer corpora lutea in their ovaries compared to control littermates, further supporting a role of MC4R in regulating ovulation. While MC4R has been documented in the ovary (11), our data of i) impaired LH surge observed in the Kiss1-MC4RKO females (Figure 2M), and ii) direct action of αMSH on Kiss1 ARC and AVPV/PeN neuronal populations through MC4R (Figures 4, 5) support a CNS role for MC4R in the regulation of fertility. However, an additional effect of MC4R at the level of the ovary cannot be ruled. One possibility is that the genetic deletion of Mc4r from Kiss1 gene may have affected the ovary. However, to test this hypothesis, the co-expression of Mc4r and Kiss1 in the ovary needs to be investigated. Another possibility is that given the partial fertility of these mice, follicles may have developed at different rates, potentially due to irregular secretion of gonadotropins. As a result, fewer follicles would develop and ovulate. In adulthood, the reproductive phenotype observed in Kiss1-MC4RKO and MC4R-loxTB female mice (increased LH, irregular estrous cycles, oligo-ovulation, increased cystic follicles) correlates with the phenotype observed in PCOS mouse models (22, 23). In fact, an association between MC4R mutations and PCOS has been reported (10). However, our mouse models failed to display higher levels of circulating androgens or AMH, two of the hallmarks of PCOS (24), suggesting that melanocortin signaling is unlikely to be a main contributing factor in the development of this syndrome.

Earlier version: Our data revealing early puberty onset and impaired LH pulse and surge driving ovulation support a direct role of MC4R on the two primary Kiss1 populations regulating these functions (ARH and AVPV/PeN). We further investigated this and found that MC4R signaling excites Kiss1ARH neurons and inhibits Kiss1AVPV/PeN neurons in an estradiol-dependent manner. In Kiss1ARH neurons, bath applied or optogenetically-evoked release of melanocortin agonists induce a direct excitatory inward current. Our studies using whole cell voltage clamp in female mice were able to detect a melanocortin-mediated inward current in Kiss1ARH neurons that would increase excitability. Previously, an MCR-activation effect was not detected in cell-attached recordings of diestrus females (14, 44). However, cell-attached recordings do not measure a direct response in isolated Kiss1ARH neurons but rather the summation of multiple synaptic inputs to Kiss1ARH neurons. Most importantly, we did not see a difference in the excitatory response in synaptically isolated Kiss1ARH neurons from E2-treated versus vehicle-treated, ovariectomized females. This emphasizes the importance of doing whole-cell recording from isolated Kiss1 neurons. The modulatory actions of POMCARH neuronal projections to Kiss1ARH neurons can also occur through a) the excitatory action of glutamate in a process that is facilitated in the presence of estradiol, in line with previous publications (33, 45), and b) the inhibitory action of β-endorphin, suggesting that POMCARH neurons are able to exert a precise regulatory role of the GnRH pulse generator, i.e. Kiss1ARH neurons, in response to metabolic challenges. In the AVPV/PeN, MC4R agonists inhibited Kiss1AVPV/PeN neurons in an estradiol-dependent manner, which utilized a Gαq-coupled membrane estrogen receptor (STX receptor) to attenuate the inhibitory tone of MC4R in these neurons. Rapid estrogen signaling may act to ease transitions between states. Membrane delimited E2 actions can quickly attenuate or enhance coupling between receptors and signaling cascades. These effects will precede E2-driven changes in gene expression that produce more stable alterations in signaling. This combination of mechanisms will reduce any lag between rises in serum E2 and physiological effects. Considering the abbreviated mouse reproductive cycle, parallel mechanisms acting on different timescales are particularly important.

Revised version: Our data revealing early puberty onset, augmented LH pulse and impaired LH surge driving ovulation support a direct role of MC4R on the two primary Kiss1 populations regulating these functions (ARH and AVPV/PeN). We further investigated this and found that MC4R signaling excites Kiss1ARH neurons and inhibits Kiss1AVPV/PeN neurons in an estradiol-dependent manner. In Kiss1ARH neurons, bath applied or optogenetically-evoked release of melanocortin agonists induce a direct excitatory inward current. Our studies using whole cell voltage clamp in female mice were able to detect a melanocortin-mediated inward current in Kiss1ARH neurons that would increase excitability. Previously, an MCR-activation effect was not detected in cell-attached recordings of diestrus females (14, 44). However, cell-attached recordings do not measure a direct response in isolated Kiss1ARH neurons but rather the summation of multiple synaptic inputs to Kiss1ARH neurons. Most importantly, we did not see a difference in the excitatory response in synaptically isolated Kiss1ARH neurons from E2-treated versus vehicle-treated, ovariectomized females. This emphasizes the importance of doing whole-cell recording from isolated Kiss1 neurons. The modulatory actions of POMCARH neuronal projections to Kiss1ARH neurons can also occur through a) the excitatory action of glutamate in a process that is facilitated in the presence of estradiol, in line with previous publications (33, 45), and b) the inhibitory action of β-endorphin, suggesting that POMCARH neurons are able to exert a precise regulatory role of the GnRH pulse generator, i.e. Kiss1ARH neurons, in response to metabolic challenges. In the AVPV/PeN, MC4R agonists inhibited Kiss1AVPV/PeN neurons in an estradiol-dependent manner, which utilized a Gαq-coupled membrane estrogen receptor (STX receptor) to attenuate the inhibitory tone of MC4R in these neurons. Rapid estrogen signaling may act to ease transitions between states. Membrane delimited E2 actions can quickly attenuate or enhance coupling between receptors and signaling cascades. These effects will precede E2-driven changes in gene expression that produce more stable alterations in signaling. This combination of mechanisms will reduce any lag between rises in serum E2 and physiological effects. Considering the abbreviated mouse reproductive cycle, parallel mechanisms acting on different timescales are particularly important.

Our pharmacology data clearly demonstrates the estradiol-dependent action of MC4R in Kiss1AVPV/PeN neurons, suggesting that similar to Kiss1ARH neurons, this population is regulated by melanocortins. The effect of MC4R agonists on Kiss1 neurons in the presence of TTX, suggests a direct synaptic effect without the need of channelrhodopsin-2(ChR2)-assisted circuit mapping. Nonetheless, to further demonstrate this interaction, we provide evidence of profound innervation of αMSH fibers and the presence of synaptic contact between POMC and Kiss1AVPV/PeN neurons through the optogenetic stimulation of a glutamatergic response in Kiss1AVPV/PeN neurons after low frequency stimulation of POMC terminals. We intentionally avoided high frequency stimulation of POMC terminals to prevent the co-release of β-endorphins (e.g., Figure 4M) and their possible additive effect on the inhibitory action of MC4R. Future studies will be aimed at the characterization of the endogenous opioid pathways in the induction of the LH surge.

To date, two instances of MC4R-mediated inhibition have been described in the central nervous system: MC4R can Gi,o couple to open K+ (Kir6.2) channels in parasympathetic preganglionic neurons in the brainstem (39), and MC4R can directly couple to K+ (Kv7.1) channels in the hypothalamic paraventricular nucleus neurons (46). Interestingly, our findings describe an inhibitory role of MC4R in the reproductive neuroendocrine axis. This steroid state dependent effect is particularly relevant because in the context of reproduction, the two populations of Kiss1 neurons are strikingly different in their response to estradiol, where peptide expression in Kiss1ARH and Kiss1AVPV/PeN neurons are inhibited or stimulated, respectively, in order to mount the negative versus positive feedback of sex steroids (15). This differential regulation allows for the episodic versus surge release of GnRH; however, the cellular mechanisms underlying these opposing roles of Kiss1 neurons to the same stimulus, i.e. circulating E2 levels, remain unknown. Although there were no differences in the magnitude of the electrophysiological responses to MTII in Kiss1ARH neurons in vehicle-versus E2-treated, ovariectomized females, we cannot rule out that E2 affects other systems activated by the MC4R signaling cascade as recently reported in MC4R expressing neurons of the ventromedial nucleus of the hypothalamus (VMH) (47). Interestingly, puberty onset was advanced in Kiss1-MC4RKO females. This data suggests that the inhibitory tone of MC4R signaling on Kiss1AVPV/PeN neurons might also be involved in the timing of puberty onset, which in turn suggests a role for this female-specific population of Kiss1 neurons in sexual maturation. Furthermore, our findings suggest that metabolic cues, through the regulation of the melanocortin output onto Kiss1AVPV/PeN neurons, are essential for the timing and magnitude of the GnRH/LH surge. This study suggests that the melanocortin-kisspeptin pathway is one of the multiple pathways essential for the metabolic regulation of the HPG axis. Future studies are warranted to fully elucidate synaptic input and postsynaptic cascades in the POMC-Kiss1AVPV/PeN circuit.

Altogether, our data reveal a differential regulatory action of MC4R in the neural control of GnRH/LH release, participating in both the surge (E2-treated, ovariectomized female) and the pulse-like (ovariectomized female) modes of LH release through the cellular regulation of Kiss1 neurons, which translate into our findings on the conditional knockout of Mc4r from Kiss1 neurons (Figure 6). These findings are important because the reproductive abnormalities often attributed to obesity in MC4R deficient patients may not be caused by the excess in body weight but, at least in part, by a deficiency in MC4R signaling directly at the level of Kiss1 neurons that affects predominantly the ability to mount a preovulatory LH surge.

Working model of the neuronal input from POMC ARH to Kiss1 ARH and AVPV populations.

POMCARH neurons excite Kiss1ARH neurons via glutamate and αMSH release, which activate glutamatergic ionotropic receptors and MC4R, respectively. In low E2 state, αMSH fibers project to the AVPV/PeN, and MC4R inhibit Kiss1AVPV/PeN neurons via opening a K+ channel. However, in high E2 state, E2 enhances the glutamatergic excitation of both POMC and Kiss1 ARH neurons. E2-treatment rapidly abrogates (uncouples) the inhibitory MC4R signaling in Kiss1AVPV/PeN neurons, allowing greater excitatory actions of Kiss1ARH neurons via enhanced glutamate release (26).

Materials and methods

In-vivo experimental procedures

Generation of Kiss1-MC4RKO, Kiss1-cre:MC4R-LoxTB and MC4R-LoxTB Transgenic Mice

Earlier version: Kiss1-MC4RKO mice were generated by crossing Kiss1-cre knock-in mice and MC4Rlox/lox mice. Kiss1-cre mice (RRID:MGI:6278139) were obtained from Dr. Richard Palmiter (University of Washington, Seattle, WA) (49) and Mc4rlox/lox (RRID:IMSR_JAX:023720) were a gift from Dr. Brad Lowell (Beth Israel Deaconess Hospital, Boston, MA) (4). These mice were crossed to generate Kiss1-MC4RKO mice lacking Mc4r expression selectively from Kiss1 neurons (Kiss1-cre+/-; MC4Rlox/lox) and their control littermates (MC4Rlox/lox). To generate Kiss1-cre:MC4R-LoxTB and MC4R-LoxTB mice, MC4R-LoxTB mice were purchased from The Jackson Laboratory (MC4R-LoxTB; catalog no. 006414). These mice present with a loxp-flanked transcriptional blocking (LoxTB) sequence preventing normal endogenous gene transcription and translation from the endogenous locus. Homozygous MC4R-loxTB/loxTB mice are devoid of functional MC4R mRNA (MC4RKO mice), while the presence of Cre recombinase on the Kiss1 promotor in Kiss1-cre mice will result in the removal of the transcription blocker and subsequent expression of MC4R in tissue-specific sites (i.e. Kiss1 neurons) therefore resulting in the generation of Kiss1-cre:MC4R-LoxTB (with Mc4r expression restored in Kiss1 neurons), their obese control littermates MC4R-LoxTB mice (MC4RKO), and their WT controls (MC4R+/+). To rule out early embryonic recombination of the MC4RloxTB/loxTB or MC4Rlox/lox alleles, we ran PCR assays on tail DNA designed to detect wild-type allele (MC4R+/+), undeleted lox (MC4Rlox/lox) or loxTB alleles (MC4RloxTB/loxTB). Genotyping was confirmed by sending tissue to Transnetyx, Inc., for testing by real-time polymerase chain reaction. Mice were housed in Harvard Medical School Animal Resources facilities where they were fed standard mouse chow (Teklad F6 Rodent Diet 8664) and were given ad libitum access to tap water under constant conditions of temperature (22– 24°C) and light (12 hr light [07:00]/dark [19:00] cycle).

Revised version: Kiss1-MC4RKO mice were generated by crossing Kiss1-cre knock-in mice and MC4Rlox/lox mice. Kiss1-cre mice (RRID:MGI:6278139) were obtained from Dr. Richard Palmiter (University of Washington, Seattle, WA) (48) and Mc4rlox/lox (RRID:IMSR_JAX:023720) were a gift from Dr. Brad Lowell (Beth Israel Deaconess Hospital, Boston, MA) (4). These mice were crossed to generate Kiss1-MC4RKO mice lacking Mc4r expression selectively from Kiss1 neurons (Kiss1-cre+/-; MC4Rlox/lox) and their control littermates (MC4Rlox/lox). To generate Kiss1-cre:MC4R-LoxTB and MC4R-LoxTB mice, MC4R-LoxTB mice were purchased from The Jackson Laboratory (MC4R-LoxTB; catalog no. 006414). These mice present with a loxp-flanked transcriptional blocking (LoxTB) sequence preventing normal endogenous gene transcription and translation from the endogenous locus. Homozygous MC4R-loxTB/loxTB mice are devoid of functional MC4R mRNA (MC4RKO mice), while the presence of Cre recombinase on the Kiss1 promotor in Kiss1-cre mice will result in the removal of the transcription blocker and subsequent expression of MC4R in tissue-specific sites (i.e. Kiss1 neurons) therefore resulting in the generation of Kiss1-cre:MC4R-LoxTB (with Mc4r expression restored in Kiss1 neurons), their obese control littermates MC4R-LoxTB mice (MC4RKO), and their WT controls (MC4R+/+). To rule out early embryonic recombination of the MC4RloxTB/loxTB or MC4Rlox/lox alleles, we ran PCR assays on tail DNA designed to detect wild-type allele (MC4R+/+), undeleted lox (MC4Rlox/lox) or loxTB alleles (MC4RloxTB/loxTB). Genotyping was confirmed by sending tissue to Transnetyx, Inc., for testing by real-time polymerase chain reaction. Mice were housed in Harvard Medical School Animal Resources facilities where they were fed standard mouse chow (Teklad F6 Rodent Diet 8664) and were given ad libitum access to tap water under constant conditions of temperature (22–24°C) and light (12 hr light [07:00]/dark [19:00] cycle).

RNAscope in situ hybridization

To validate the Kiss1-MC4RKO, Kiss1-cre:MC4R-LoxTB and MC4R-LoxTB (MC4RKO) mouse models and investigate the co-expression of Kiss1 and Mc4r mRNA in these mice, dual fluorescence ISH was performed using RNAscope (ACD, Multiplex Fluorescent v.2) according to the manufacturer’s protocol using the following probes: Mc4r (319181-C2) and Kiss1 (500141-C1). Brains (n=4/group) from WT OVX (for expression in the ARH) and OVX+E2 (for expression in the AVPV/PeN) mice and OVX Kiss1-MC4RKO, Kiss1-cre:MC4R-LoxTB and MC4RKO were removed fresh frozen on dry ice, and then stored at -80°C until sectioned. Five sets of 20 μm sections in the coronal plane were cut on a cryostat, from the diagonal band of Broca to the mammillary bodies, thaw mounted onto SuperFrost Plus slides (VWR Scientific) and stored at -80°C until use. A single set was used for ISH experiment (adjacent sections 100 mm apart). Images were taken at 20x magnification of the sections containing AVPV, PeN, and the three rostro-to-caudal levels of the ARH, and Kiss1 neurons expressing (Kiss1-cre:MC4R-LoxTB mice) or lacking (Kiss1-MC4RKO mice) Mc4r were identified using ImageJ.

Reproductive maturation of Kiss1-MC4RKO, Kiss1-cre:MC4R-LoxTB and MC4R-LoxTB mice

To assess the reproductive phenotype of mice with selective reinsertion of Mc4r in Kiss1 neurons (Kiss1-cre:MC4R-LoxTB, n=11), selective deletion of Mc4r from Kiss1 neurons (Kiss1-MC4RKO, n=24), global deletion of Mc4r (MC4R-LoxTB, n=7); and their control MC4Rlox/lox littermates (n=14), mice were weaned at post-natal day (PND) 21 and were monitored daily for puberty onset. Females were monitored daily for vaginal opening (VO, indicative of the complete canalization of the vaginal cavity) and for first estrus (first day with cornified cells determined by daily morning vaginal cytology) after the day of VO. Body weight (BW) was measured at the day of puberty onset to determine if changes in puberty onset could be due to differences in BW. Estrous cyclicity was monitored in females by daily vaginal cytology, for a period of 15 days, in 6-month-old mice and their respective control littermates (n=5/group). Cytology samples were obtained every morning (9 am), placed on a glass slide and stained with hematoxylin and eosin for determination of the estrous cycle stage under the microscope.

Fecundity test of Kiss1-MC4RKO females

Adult 6-month-old Kiss1-MC4RKO and control littermate female mice (n=9/group) were placed with adult WT males proven to father litters for 90 days and time to deliver pups and number of pups per litter were monitored.

Characterization of the estradiol-induced luteinizing hormone surge

Earlier version: Kiss1-MC4RKO (n=6) and control MC4Rlox/lox littermate (n=5) adult female mice were subjected to bilateral ovariectomy (OVX) via abdominal incision under light isoflurane anesthesia. Immediately after OVX, capsules filled with E2 (1 ug/20g BW) were implanted subcutaneously via a small mid-scapular incision on the back. Five days later, mice were subcutaneously injected in the morning with estradiol benzoate (1 ug/20g BW) to produce elevated proestrus-like E2 levels (preovulatory LH surge) on the following day (50). Blood samples were collected at 8am and 7pm; LH levels were stored at -80°C until measured via LH ELISA.

Revised version: Kiss1-MC4RKO (n=6) and control MC4Rlox/lox littermate (n=5) adult female mice were subjected to bilateral ovariectomy (OVX) via abdominal incision under light isoflurane anesthesia. Immediately after OVX, capsules filled with E2 (1 ug/20g BW) were implanted subcutaneously via a small mid-scapular incision on the back. Five days later, mice were subcutaneously injected in the morning with estradiol benzoate (1 ug/20g BW) to produce elevated proestrus-like E2 levels (induced LH surge) on the following day (49). Blood samples were collected at 8am and 7pm; LH levels were stored at -80°C until measured via LH ELISA.

Ovarian histology and hormone measurements

Bilateral ovariectomy of 6-month-old Kiss1-MC4RKO (n=4), Kiss1-cre:MC4R-LoxTB, MC4R-LoxTB and their control littermates (n=5/group) was performed under light isoflurane anesthesia. Briefly, the ventral skin was shaved and cleaned, and one small abdominal incision was made. Once the ovaries were identified and excised, the muscle incision was sutured, and the skin was closed with surgical clips. Ovaries were stored in Bouin’s fixative, sectioned, and stained with hematoxylin and eosin at the Harvard Histopathology Core. Corporal lutea were counted in the middle section of each ovary. Serum samples were also collected for analysis of testosterone and anti-müllerian hormone (AMH) levels in these mice. These hormones levels were measured at the University of Virginia Ligand Assay core with the Mouse & Rat Testosterone ELISA assay (reportable average range 10-1600 ng/dL; sensitivity of 10 ng/dL); AMH ELISA assay (reportable average range 0.2 – 15 ng/mL; sensitivity of 0.2 ng/mL).

LH pulsatile secretion profile in gonad intact Kiss1-MC4RKO, Kiss1-cre:MC4R-LoxTB and MC4R-LoxTB female mice

Earlier version: To assess the profile of LH pulses secretion, adult 6 months old gonad intact Kiss1-MC4RKO females (n=6), their control Mc4rlox/lox littermates (n=7); and Kiss1-cre:MC4R-LoxTB (n=5/group), MC4R-LoxTB and their control WT females (n=4/group) were handled daily for 3 weeks to allow acclimation to sampling conditions prior to the experiment. Pulsatile measurements of LH secretion were assessed in diestrus by repeated blood collection through a single incision at the tip of the tail. The tail was cleaned with saline and 4 μL of blood was taken at each time point from the cut tail with a pipette. We collected sequential blood samples every 10 minutes over a 180-minute sampling period. Samples were immediately frozen on dry ice and stored at −80 °C until analyzed with LH ELISA as previously described (51). The functional sensitivity of the ELISA assay was 0.0039 ng/ml with a CV% of 3.3%.

Revised version: To assess the profile of LH pulses secretion, adult 6 months old gonad intact Kiss1-MC4RKO females (n=6), their control Mc4rlox/lox littermates (n=7); and Kiss1-cre:MC4R-LoxTB (n=5/group), MC4R-LoxTB and their control WT females (n=4/group) were handled daily for 3 weeks to allow acclimation to sampling conditions prior to the experiment. Pulsatile measurements of LH secretion were assessed in diestrus by repeated blood collection through a single incision at the tip of the tail. The tail was cleaned with saline and 4 μL of blood was taken at each time point from the cut tail with a pipette. We collected sequential blood samples every 10 minutes over a 180-minute sampling period. Samples were immediately frozen on dry ice and stored at −80 °C until analyzed with LH ELISA as previously described (50). The functional sensitivity of the ELISA assay was 0.0039 ng/ml with a CV% of 3.3%.

LH pulses analysis

Earlier version: LH pulses in mice were analyzed using a custom-made MATLAB-based algorithm. The MATLAB code includes a loop that determines LH pulses as any LH peak: (i) whose height is 20% greater than the heights of the 2 previous values; (ii) 10% greater than the height of the following value; and (iii) the peak at the second time interval needs to be 20% greater than the single value that comes before it to be considered a pulse, as we previously described (52).

Revised version: LH pulses in mice were analyzed using a custom-made MATLAB-based algorithm. The MATLAB code includes a loop that determines LH pulses as any LH peak: (i) whose height is 20% greater than the heights of the 2 previous values; (ii) 10% greater than the height of the following value; and (iii) the peak at the second time interval needs to be 20% greater than the single value that comes before it to be considered a pulse, as we previously described (51).

LH pulsatility was assessed by measuring: (1) the total secretory mass, assessed by area under the curve (AUC); (2) the LH pulse amplitude, calculated by averaging the 4 highest LH values in the samples collection period for each animal; (3) the basal LH, calculated by averaging the 4 lowest LH values in the samples collection period for each animal; and (4) the total number of pulses throughout the 180 minutes sampling period.

Immunohistochemistry

Animals and Treatment

Coronal brain blocks (2 mm each) from adult female C57BL/6 mice (n=4) were fixed by immersion in 4% paraformaldehyde for ∼ 8 hours, cryoprotected in 30% sucrose solution, frozen in isopentane at –55°C, sectioned coronally on a cryostat at 20 μm, and thaw-mounted onto Superfrost Plus slides (Fisher Scientific, Pittsburgh, PA). The 20 µm sections were stored at –20°C until used for immunocytochemistry.

Immunocytochemistry

Earlier version: The sections were rinsed in PB (0.1M phosphate buffer, pH 7.4) for at least 30 min. Next, sections were incubated with normal serum corresponding to the host for the secondary antiserum (5 % normal serum with 0.3% Triton-X100 in PBS for 30 min), rinsed in PB and then incubated for ∼45 h at 4° C with a rabbit polyclonal antiserum against α-MSH (1:2,500). The specificity of this antiserum has been documented (53). After rinsing, sections were first incubated for 2-3 hours at room temperature with biotinylated donkey anti rabbit gamma globulin (IgG; 1:500) and next with streptavidin-Alexa 488 (1:2500). Both the primary and secondary antibodies were diluted in tris-(hydroxymethyl) aminomethane (0.5 %, Jackson ImmunoResearch, Philadelphia, PA) in PB containing 0.7 % seaweed gelatin (Jackson ImmunoResearch, Philadelphia, PA) and 0.5 % Triton X-100 and 3% bovine serum albumin (BSA; Jackson ImmunoResearch, Philadelphia, PA) adjusted to pH 7.6. Following a final rinse overnight, slides were coverslipped with gelvatol containing the anti-fading agent, 1,4-diazabicyclo(2,2)octane (DABCO; Cold Spring Harbor Protocols, 2006).

Revised version: The sections were rinsed in PB (0.1M phosphate buffer, pH 7.4) for at least 30 min. Next, sections were incubated with normal serum corresponding to the host for the secondary antiserum (5 % normal serum with 0.3% Triton-X100 in PBS for 30 min), rinsed in PB and then incubated for ∼45 h at 4° C with a rabbit polyclonal antiserum against α-MSH (1:2,500). The specificity of this antiserum has been documented (52). After rinsing, sections were first incubated for 2-3 hours at room temperature with biotinylated donkey anti rabbit gamma globulin (IgG; 1:500) and next with streptavidin-Alexa 488 (1:2500). Both the primary and secondary antibodies were diluted in tris-(hydroxymethyl) aminomethane (0.5 %, Jackson ImmunoResearch, Philadelphia, PA) in PB containing 0.7 % seaweed gelatin (Jackson ImmunoResearch, Philadelphia, PA) and 0.5 % Triton X-100 and 3% bovine serum albumin (BSA; Jackson ImmunoResearch, Philadelphia, PA) adjusted to pH 7.6. Following a final rinse overnight, slides were coverslipped with gelvatol containing the anti-fading agent, 1,4-diazabicyclo(2,2)octane (DABCO; Cold Spring Harbor Protocols, 2006).

Imaging

Photomicrographs of labeling were initially acquired using a Nikon E800 fluorescent microscope (Eclipse E800; Nikon Instruments, Melville, NY) equipped with a fiber illuminator (Intensilight C-HGFI; Nikon Instruments) and a high-definition digital microscope camera head (DS-Fi1; Nikon Instruments) interfaced with a PC-based camera controller (DS-U3; Nikon Instruments).

Real time quantitative PCR

Earlier version: i) To investigate the changes in the expression of the melanocortin genes Agrp, Pomc, Mc3r and Mc4r in the MBH during development in prepubertal and pubertal WT females at ages P10 (n=6), P15 (n=6), P22 (n=6), and P30 (n=5), and ii) evaluate the gene expression profile of Pdyn, Kiss1, Tac2, Tacr3, in the ARH of adult intact (in diestrus) female Kiss1-MC4RKO (n=6) and their control MC4Rlox/lox littermates (n=3), the brains were removed and rapidly embedded in Tissue-Tek, frozen in −30°C isopentane solution and stored at −80°C until use. Frozen tissue punches were recovered through MBH with a 1 mm diameter canula (54). These tissue punches encompassed the whole MBH from the WT females (figure 1D) and the ARH from the Kiss1-MC4RKO and their control females (Figure 2I). The tissues were homogenized, and total RNA was isolated using TRIzol reagent (Invitrogen) followed by chloroform/isopropanol extraction. RNA purity and concentration were measured via an absorbance spectrophotometer (260/280 nm > 1.8; NanoDrop 1000, Thermo Fisher Scientific). Total RNA (1 μg) was reverse transcribed to cDNA using random hexamers (High-Capacity cDNA Synthesis Kit, Life Technologies). Quantitative real-time PCR assays were performed using SYBR Green RT-qPCR master mix (Applied Biosystems) and analyzed using ABI Prism 7000 SDS software (Applied Biosystems). The cycling conditions were: 2 minutes incubation at 95°C (hot start), 45 amplification cycles (95°C for 30 seconds, 60°C for 30 seconds, and 45 seconds at 75°C, with fluorescence detection at the end of each cycle), followed by melting curve of the amplified products obtained by ramped increase of the temperature from 55°C to 95°C to confirm the presence of single amplification product per reaction. PCR specificity was verified by melting curve analysis and agarose gel electrophoresis. Each sample was run in duplicate to obtain an average cycle threshold (CT) value and relative expression of each target gene was determined using the comparative Ct method (55). The data were normalized to Hypoxanthine Guanine Phosphoribosyltransferase (Hprt) expression levels in each sample. Results were expressed as fold differences in relative gene expression with respect to i) P10 for melanocortin genes expression analysis during development in WT mice, and ii) controls for the KNDy genes expression in female Kiss1-MC4RKO mice. The primers used are listed in Table 1.

Revised version: i) To investigate the changes in the expression of the melanocortin genes Agrp, Pomc, Mc3r and Mc4r in the MBH during development in prepubertal and pubertal WT females at ages P10 (n=6), P15 (n=6), P22 (n=6), and P30 (n=5), and ii) evaluate the gene expression profile of Pdyn, Kiss1, Tac2, Tacr3, in the ARH of adult intact (in diestrus) female Kiss1-MC4RKO (n=6) and their control MC4Rlox/lox littermates (n=3), the brains were removed and rapidly embedded in Tissue-Tek, frozen in −30°C isopentane solution and stored at −80°C until use. Frozen tissue punches were recovered through MBH with a 1 mm diameter canula (53). These tissue punches encompassed the whole MBH from the WT females (figure 1D) and the ARH from the Kiss1-MC4RKO and their control females (Figure 2I). The tissues were homogenized, and total RNA was isolated using TRIzol reagent (Invitrogen) followed by chloroform/isopropanol extraction. RNA purity and concentration were measured via an absorbance spectrophotometer (260/280 nm > 1.8; NanoDrop 1000, Thermo Fisher Scientific). Total RNA (1 μg) was reverse transcribed to cDNA using random hexamers (High-Capacity cDNA Synthesis Kit, Life Technologies). Quantitative real-time PCR assays were performed using SYBR Green RT-qPCR master mix (Applied Biosystems) and analyzed using ABI Prism 7000 SDS software (Applied Biosystems). The cycling conditions were: 2 minutes incubation at 95°C (hot start), 45 amplification cycles (95°C for 30 seconds, 60°C for 30 seconds, and 45 seconds at 75°C, with fluorescence detection at the end of each cycle), followed by melting curve of the amplified products obtained by ramped increase of the temperature from 55°C to 95°C to confirm the presence of single amplification product per reaction. PCR specificity was verified by melting curve analysis and agarose gel electrophoresis. Each sample was run in duplicate to obtain an average cycle threshold (CT) value and relative expression of each target gene was determined using the comparative Ct method (54). The data were normalized to Hypoxanthine Guanine Phosphoribosyltransferase (Hprt) expression levels in each sample. Results were expressed as fold differences in relative gene expression with respect to i) P10 for melanocortin genes expression analysis during development in WT mice, and ii) controls for the KNDy genes expression in female Kiss1-MC4RKO mice. The primers used are listed in Table 1.

Data analysis

Statistical data are expressed as means ± SEM, where n represents the number of animals in each study group. The significance of differences between groups was evaluated using unpaired two-tailed Student’s t test, or a one- or two-way ANOVA test (with post hoc comparisons). Significance level was set at p<0.05. All analyses were performed with GraphPad Prism Software, Inc (San Diego, CA).

In-vitro experimental procedures

Animals

Earlier version: All animal procedures described in the electrophysiology studies were performed in accordance with institutional guidelines based on National Institutes of Health standards and approved by the Institutional Animal Care and Use Committees at Oregon Health and Science University and Appalachian State University. Kiss1-cre (v2) mice (49) were crossed with Ai32 (56) or C57B6J mice. Pomc-Cre mice (RRID:IMSR_JAX:005965) (57) were crossed with wildtype C57B6J (RRID:IMSR_JAX:000664) mice. The Ai32 cross was not used with Pomc-Cre mice because the gene is transiently expressed during development in some cells fated to be Kiss1 or AgRP cells (58) and early recombinant events lead to persistent expression of the ChR2-mCh fusion protein in non-POMC cells. However, we have previously shown that AAV-driven expression in adult Pomc-cre animals is restricted to β-endorphin labeled cells in the ARH (45), avoiding this non-specificity issue. All colonies were maintained onsite under controlled temperature (21-23 °C) and photoperiod (12:12-h light-dark cycle 0600 to 1800) while receiving ad libitum food (5L0D; LabDiet, St. Louis, MO) and water access. Following surgeries, mice received a s.c. dose of 4-5 mg/kg carprofen (Rimadyl; Pfizer Animal Health, New York, NY) and given at least 1 week of recovery.

Revised version: All animal procedures described in the electrophysiology studies were performed in accordance with institutional guidelines based on National Institutes of Health standards and approved by the Institutional Animal Care and Use Committees at Oregon Health and Science University and Appalachian State University. Kiss1-cre (v2) mice (48) were crossed with Ai32 (55) or C57B6J mice. Pomc-Cre mice (RRID:IMSR_JAX:005965) (56) were crossed with wildtype C57B6J (RRID:IMSR_JAX:000664) mice. The Ai32 cross was not used with Pomc-Cre mice because the gene is transiently expressed during development in some cells fated to be Kiss1 or AgRP cells (57) and early recombinant events lead to persistent expression of the ChR2-mCh fusion protein in non-POMC cells. However, we have previously shown that AAV-driven expression in adult Pomc-cre animals is restricted to β-endorphin labeled cells in the ARH (45), avoiding this non-specificity issue. All colonies were maintained onsite under controlled temperature (21-23 °C) and photoperiod (12:12-h light-dark cycle 0600 to 1800) while receiving ad libitum food (5L0D; LabDiet, St. Louis, MO) and water access. Following surgeries, mice received a s.c. dose of 4-5 mg/kg carprofen (Rimadyl; Pfizer Animal Health, New York, NY) and given at least 1 week of recovery.

Ovariectomies

Earlier version: Seven to ten days prior to each experiment, ovaries were removed as described previously while under isoflurane anesthesia (45). Two days before experiments females received either an injection of sesame oil (50 μl, sc; Sigma-Aldrich, St. Louis, MO) or a priming dose (0.25 μg/50 μl sesame oil, sc) of E2 benzoate (Sigma-Aldrich, St. Louis, MO) in the morning. On the following day, oil or a high (1.5 μg) dose of E2 benzoate, which generates a LH surge, was administered (59). Circulating levels of E2 were verified by the uterine weights (<25 mg for OVX and >95 mg for OVX E2-treated) at the time of euthanasia (59).

Revised version: Seven to ten days prior to each experiment, ovaries were removed as described previously while under isoflurane anesthesia (45). Two days before experiments females received either an injection of sesame oil (50 μl, sc; Sigma-Aldrich, St. Louis, MO) or a priming dose (0.25 μg/50 μl sesame oil, sc) of E2 benzoate (Sigma-Aldrich, St. Louis, MO) in the morning. On the following day, oil or a high (1.5 μg) dose of E2 benzoate, which generates an induced LH surge, was administered (58). Circulating levels of E2 were verified by the uterine weights (<25 mg for OVX and =95 mg for OVX E2-treated) at the time of euthanasia (58).

AAV Delivery

Bilateral ARH injections of AAV1-Ef1a-DIO-ChR2:mCherry (RRID:Addgene_20297) or AAV1-Ef1a-DIO-ChR2:YFP (RRID:Addgene_20298) were performed on adult Kiss1-cre mice or Pomc-cre mice on a stereotaxic frame under isoflurane anesthesia. ARH injection coordinates were anteroposterior (AP): −1.10 mm, mediolateral (ML): ± 0.30 mm, dorsoventral (DL): −5.80 mm (surface of brain z = 0.0 mm); 400 nl of the AAV (2.0 x 1012 particles/ml) was injected (100 nl/min) into each position. Mice were given carprofen for analgesia and allowed to recover for at least two weeks before euthanasia.

Visualized whole-cell patch recordings

Earlier version: Hypothalamic coronal brain slices (240 um) were made from female mice, using a Leica VT1000S vibratome, in ice cold cutting solution bubbled with O2/C02 (95%/5%). Slices were then transferred to a holding chamber with artificial cerebrospinal fluid bubbled with the same gas mix and allowed to recover for at least 1 hour. For recordings, slices were placed in a perfusion chamber and visualized with an Olympus BX51W1 using either differential infrared contrast or oblique illumination. Kiss1 neurons in the ARH or AVPV/PeN were targeted for electrophysiological recordings as done previously and described below (60, 61). Final concentration was calculated based on the drug in the known volume of the bath. Perfused drugs (TTX, naloxone, and STX) were constantly circulated and given at least 10 minutes to reach maximal effect. Focal application of drugs (MTII and THIQ) was done with the pump stopped (10-20 minutes). 0.3 to 1 μl was directly added to the known volume of the bath to achieve the desired final concentration. This approach enables precise timing of recordings and reduces the risk of receptor desensitization.

Revised version: Hypothalamic coronal brain slices (240 um) were made from female mice, using a Leica VT1000S vibratome, in ice cold cutting solution bubbled with O2/C02 (95%/5%). Slices were then transferred to a holding chamber with artificial cerebrospinal fluid bubbled with the same gas mix and allowed to recover for at least 1 hour. For recordings, slices were placed in a perfusion chamber and visualized with an Olympus BX51W1 using either differential infrared contrast or oblique illumination. Kiss1 neurons in the ARH or AVPV/PeN were targeted for electrophysiological recordings as done previously and described below (59, 60). Final concentration was calculated based on the drug in the known volume of the bath. Perfused drugs (TTX, naloxone, and STX) were constantly circulated and given at least 10 minutes to reach maximal effect. Focal application of drugs (MTII and THIQ) was done with the pump stopped (10-20 minutes). 0.3 to 1 μl was directly added to the known volume of the bath to achieve the desired final concentration. This approach enables precise timing of recordings and reduces the risk of receptor desensitization.

Solutions/drugs

Earlier version: Standard vibratome slicing, external, and internal recording solutions were utilized as previously described (60, 61). Tetrodotoxin was purchased from Alomone Labs (Jerusalem, Israel), Melanotan II and αMSH from Tocris (Minneapolis, MN). TEA, 4-AP, 17β-estradiol benzoate, and Naloxone were purchased from Millipore-Sigma. STX was produced by AAPharmaSyn, LLC (Ann Arbor, MI) under contract.

Revised version: Standard vibratome slicing, external, and internal recording solutions were utilized as previously described (59, 60). Tetrodotoxin was purchased from Alomone Labs (Jerusalem, Israel), Melanotan II and αMSH from Tocris (Minneapolis, MN). TEA, 4-AP, 17β-estradiol benzoate, and Naloxone were purchased from Millipore-Sigma. STX was produced by AAPharmaSyn, LLC (Ann Arbor, MI) under contract.

Electrophysiology data analysis

Electrophysiological data were analyzed using Clampfit 10/11 (Molecular Devices) and Prism 7/10 (Dotmatics). All values are expressed as Mean ± SEM. Comparisons between two groups were made using un-paired Student’s t-test or between multiple groups using an ANOVA (with post hoc comparisons) with p-values < 0.05 considered significant. When variances differed significantly, Mann-Whitney U test was used instead.

Targeting of Kiss1 neurons for electrophysiological recordings

Earlier version: For non-optogenetic experiments, brain slices were taken from AAV injected Kiss1-Cre AAV or Kiss1xAi32 female mice. Ai32 mice (RRID:IMSR_JAX:024109, C57BL/6 background) carry the floxed ChR2 (H134R)-EYFP gene in their Gt(ROSA)26Sor Locus (56), allowing its expression in a Cre-dependent manner. Due to concerns of nonspecific expression (62), we previously validated this model using single cell RT-PCR and documented that Kiss1 mRNA was detectable in 99% of individually harvested eYFP cells (n=126). In addition, we have used both AAV injected Kiss1-Cre AAV or Kiss1xAi32 female mice and found no differences in electrophysiological results (60). We used AAV injected POMC-Cre female mice and avoided small soma, high input resistance (>800 MΩ), ventrally located cells that are typically NPY/AgRP neurons. Instead, we targeted larger, more dorsomedial neurons in the ARH while avoiding fluorescent (POMC) cells. Unlike ARH Kiss1 neurons, POMC neurons do not make reciprocal projections, so uninfected POMC neurons do not display monosynaptic EPSCs in response to optogenetic stimulation. Next, we used a ramp IV protocol to probe for the presence of a persistent sodium current (INaP). While this current is more prevalent in the AVPV Kiss1 population, Kiss1 neurons are the only ARH neurons to display this electrophysiological “fingerprint” of a pronounced INaP paired with a high capacitance and low input resistance (26). Finally, we also harvested the cytosol at the end of all recordings and measured the expression of Kiss1. Only cells that expressed a persistent sodium current and/or expressed Kiss1 mRNA were included in final analysis (n = 40 ARH & 6 AVPV).

Revised version: For non-optogenetic experiments, brain slices were taken from AAV injected Kiss1-Cre AAV or Kiss1xAi32 female mice. Ai32 mice (RRID:IMSR_JAX:024109, C57BL/6 background) carry the floxed ChR2 (H134R)-EYFP gene in their Gt(ROSA)26Sor Locus (55), allowing its expression in a Cre-dependent manner. Due to concerns of nonspecific expression (61), we previously validated this model using single cell RT-PCR and documented that Kiss1 mRNA was detectable in 99% of individually harvested eYFP cells (n=126). In addition, we have used both AAV injected Kiss1-Cre AAV or Kiss1xAi32 female mice and found no differences in electrophysiological results (59). We used AAV injected POMC-Cre female mice and avoided small soma, high input resistance (>800 MΩ), ventrally located cells that are typically NPY/AgRP neurons. Instead, we targeted larger, more dorsomedial neurons in the ARH while avoiding fluorescent (POMC) cells. Unlike ARH Kiss1 neurons, POMC neurons do not make reciprocal projections, so uninfected POMC neurons do not display monosynaptic EPSCs in response to optogenetic stimulation. Next, we used a ramp IV protocol to probe for the presence of a persistent sodium current (INaP). While this current is more prevalent in the AVPV Kiss1 population, Kiss1 neurons are the only ARH neurons to display this electrophysiological “fingerprint” of a pronounced INaP paired with a high capacitance and low input resistance (26). Finally, we also harvested the cytosol at the end of all recordings and measured the expression of Kiss1. Only cells that expressed a persistent sodium current and/or expressed Kiss1 mRNA were included in final analysis (n = 40 ARH & 6 AVPV).